Acquisition of auditory discrimination mediated by different processes through two distinct circuits linked to the lateral striatum

Author response:

The following is the authors’ response to the previous reviews.

Although we have no further revisions on the manuscript, we would like to respond to the remaining comments from the reviewers as follows.

Reviewer 1:

The authors have addressed some concerns raised in the initial review but some remain. In particular it is still unclear what conclusions can be drawn about taskrelated activity from scans that are performed 30 minutes after the behavioral task. I continue to think that a reorganization/analysis data according to event type would be useful and easier to interpret across the two brain areas, but the authors did not choose to do this. Finally, switching the cue-response association, I am convinced, would help to strengthen this study.

As for the task-related activity, the strategy for PET scan was explained in our response to the comment 2 from Reviewer 2. Briefly, rats receive intravenous administration of 18F-FDG solution before the start of the behavioral session. The 18FFDG uptake into the cells starts immediately and reaches the maximum level until 30 min, being kept at least for 1 h. A 30-min PET scan is executed 25 min after the session. Therefore, the brain activity reflects the metabolic state during task performance in rats.

Regarding data presentation of the electrophysiological experiments, we described the subpopulations of event-related neurons showing notable neuronal activity patterns in the order of aDLS and pVLS, according to the procedure of explanations for the behavioral study

For switching the cue-response association, we mentioned the difference in firing activity between HR and LL trials, suggesting that different combinations between the stimulus and response may affect the level of firing activity. As suggested by the reviewer, an examination of switching the cue-response association is useful to confirm our interpretation. We will address this issue in our future studies.

Reviewer 2:

The authors have made important revisions to the manuscript and it has improved in clarity. They also added several figures in the rebuttal letter to answer questions by the reviewers. I would ask that these figures are also made public as part of the authors' response or if not, included in the manuscript.

We will present the figures publicly available as part of our response.

Changes from the earlier version to the revised version

Abstract

Abstract

The striatum, the central hub of cortico-basal ganglia loops, contains functionally heterogeneous subregions distinguished by the topographic patterns of structural connectivity. These subregions mediate various processes of procedural learning. However, it remains unclear when and how striatal subregions engage in the acquisition of sensory stimulus-based decision-making. A neuroimaging of regional brain activity shows that the anterior dorsolateral striatum (aDLS) and posterior ventrolateral striatum (pVLS) in rats are activated in a different temporal pattern during the acquisition phase of auditory discrimination. Chronic and transient pharmacologic manipulations show that the aDLS promotes the behavioral strategy driven by the stimulus- response association while suppressing that by the response-outcome association, and that the pVLS contributes to forming and maintaining the stimulus-response strategy. Electrophysiological recording indicates that subpopulations of aDLS neurons predominantly represent the outcome of specific behaviors at the initial period of discrimination learning, and that pVLS subpopulations encode the beginning and ending of each behavior according to the progress of learning. In addition, other subpopulations of striatal neurons indicate sustained activation after obtaining reward with distinct patterns reflecting the stimulus-response associations. Our findings demonstrate that aDLS and pVLS neurons integrate the new learning of auditory discrimination in spatiotemporally and functionally different manners.

Introduction

Earlier version: The neural network connecting the cerebral cortex, basal ganglia, and thalamus forms functionally heterogeneous loops distinguished by the topographic patterns of connectivity1,2, and is implicated in various learning processes by the integration of cognitive, sensorimotor, and reward information3,4. The striatum is a key node in the cortico-basal ganglia loops, and its segregated subregions are considered to mediate procedural learning, such as motor skills and instrumental behavior, through parallel processing via different loops5–7. Previous studies in humans and non-human primates have reported that the neural activity within the striatum shifts from the associative subregions (caudate and anterior putamen) to the sensorimotor subregions (posterior putamen) as motor learning progresses8–11. In rodents, an instrumental behavior initiates as goal-directed actions and then transits to stimulus-response habits12,13, requiring the dorsomedial striatum (DMS; a homologue to the caudate in primates) and the dorsolateral striatum (DLS; a homologue to the putamen in primates), respectively14,15. Based on these results, a model of procedural learning has been proposed, in which the functional dominance changes from the associative to sensorimotor subregions in the striatum during the learning phases16.

Revised version: The neural network connecting the cerebral cortex, basal ganglia, and thalamus forms functionally heterogeneous loops distinguished by the topographic patterns of connectivity (Alexander and Crutcher, 1990; Foster et al., 2021), and is implicated in various learning processes by the integration of cognitive, sensorimotor, and reward information (Kim and Hikosaka, 2015; Yin and Knowlton, 2006). The striatum is a key node in the cortico-basal ganglia loops, and its segregated subregions are considered to mediate procedural learning, such as motor skills and instrumental behavior, through parallel processing via different loops (Kupferschmidt et al., 2017; Thorn et al., 2010; Yin et al., 2009). Previous studies in humans and non-human primates have reported that the neural activity within the striatum shifts from the associative subregions (caudate and anterior putamen) to the sensorimotor subregions (posterior putamen) as motor learning progresses (Doyon et al., 2009; Lehéricy et al., 2005; Miyachi et al., 2002; Miyachi et al., 1997). In rodents, an instrumental behavior initiates as goal-directed actions and then transits to stimulus-response habits (Dickinson, 1995; O’Hare et al., 2016), requiring the dorsomedial striatum (DMS; a homologue to the caudate in primates) and the dorsolateral striatum (DLS; a homologue to the putamen in primates), respectively (Yin et al., 2004; Yin et al., 2005). Based on these results, a model of procedural learning has been proposed, in which the functional dominance changes from the associative to sensorimotor subregions in the striatum during the learning phases (Redgrave et al., 2010).

Earlier version: In contrast, deviations from the prior model are evident in other behavioral tasks on decision-making based on external sensory cues. For example, the neural activity in the caudate in humans is associated with learning of a visuomotor association task and the execution of learned behavior17,18. The tail of the caudate in non-human primates contributes to the automatic, habitual behavior in decision-making task using complex visual stimuli19. In rodents, the DLS is involved in the acquisiton and execution of the discrimination task with visual and auditory cues, wheras the DMS madiates the performance of learned behavior20,21. Moreover, other striatal subregions, such as the posterior striatum and ventrolateral striatum in rodents are also engaged in the auditory discrimination22–24. Based on these results, a hypothesis has arisen that the acquisition of sensory cue-based decision-making requires wide-range spatiotemporal processing within the striatum. However, the neural mechanism by which the functional circuits linked to striatal subregions change throughout learning processes is not yet fully understood.

Revised version: In contrast, deviations from the prior model are evident in other behavioral tasks on decision-making based on external sensory cues. For example, the neural activity in the caudate in humans is associated with learning of a visuomotor association task and the execution of learned behavior (Seger and Cincotta, 2005; Toni and Passingham, 1999). The tail of the caudate in non- human primates contributes to the automatic, habitual behavior in decision-making task using complex visual stimuli (Kim and Hikosaka, 2013). In rodents, the DLS is involved in the acquisiton and execution of the discrimination task with visual and auditory cues, wheras the DMS madiates the performance of learned behavior (Featherstone and McDonald, 2004; Featherstone and McDonald, 2005). Moreover, other striatal subregions, such as the posterior striatum and ventrolateral striatum in rodents are also engaged in the auditory discrimination (Guo et al., 2019; Lee et al., 2020; Xiong et al., 2015). Based on these results, a hypothesis has arisen that the acquisition of sensory cue-based decision-making requires wide-range spatiotemporal processing within the striatum. However, the neural mechanism by which the functional circuits linked to striatal subregions change throughout learning processes is not yet fully understood.

Earlier version: To investigate a large-scale striatal reorganization in accordance with the acquisition of external cue-dependent decisions, we conducted a two-alternative auditory discrimination task with rats25. Using a small-animal neuroimaging technique, we found that the anterior dorsolateral striatum (aDLS) and posterior ventrolateral striatum (pVLS) are activated to the highest level at the middle and late stages during the acquisition phase of auditory discrimination, respectively. Then, chronic pharmacologic manipulations confirmed that the aDLS and pVLS are necessary for discrimination learning, and the transient manipulations also indicated the function of these subregions in the learning at the corresponding stages. The resultant effects of the manipulations on the behavioral strategy showed that the aDLS promotes the strategy based on the stimulus-response association, while suppressing that based on response-outcome associations. Additionally, the pVLS engages in the formation and maintenance of the stimulus-response association strategy.

Revised version: To investigate a large-scale striatal reorganization in accordance with the acquisition of external cue-dependent decisions, we conducted a two-alternative auditory discrimination task with rats (Nishizawa et al., 2012). Using a small-animal neuroimaging technique, we found that the anterior dorsolateral striatum (aDLS) and posterior ventrolateral striatum (pVLS) are activated to the highest level at the middle and late stages during the acquisition phase of auditory discrimination, respectively. Then, chronic pharmacologic manipulations confirmed that the aDLS and pVLS are necessary for discrimination learning, and the transient manipulations also indicated the function of these subregions in the learning at the corresponding stages. The resultant effects of the manipulations on the behavioral strategy showed that the aDLS promotes the strategy based on the stimulus-response association, while suppressing that based on response-outcome associations. Additionally, the pVLS engages in the formation and maintenance of the stimulus-response association strategy. Electrophysiological data of striatal neurons indicated the presence of aDLS subpopulations, mainly representing the outcome of specific behaviors at the initial period of auditory discrimination and pVLS subpopulations encoding the beginning and ending of behavior in association with the development of the discrimination. There were other subpopulations in the aDLS and pVLS showing long-term increased activity following reward with different patterns of the combination between the stimulus and response. These findings revealed that the two striatal subregions are involved in the acquisition of auditory discrimination through distinct spatiotemporal and functional fashions, challenging the prior model which proposed a transition from the associative to sensorimotor striatal subregions during learning.

Removed in the revised version: Electrophysiological data of striatal neurons indicated the presence of aDLS subpopulations, mainly representing the outcome of specific behaviors at the initial period of auditory discrimination and pVLS subpopulations encoding the beginning and ending of behavior in association with the development of the discrimination. There were other subpopulations in the aDLS and pVLS showing long-term increased activity following reward with different patterns of the combination between the stimulus and response. These findings revealed that the two striatal subregions are involved in the acquisition of auditory discrimination through distinct spatiotemporal and functional fashions, challenging the prior model which proposed a transition from the associative to sensorimotor striatal subregions during learning.

Results

Leaning processes of auditory discrimination

Earlier version: We employed a two-alternative auditory discrimination task, in which rats are required to make new associations between a tone instruction cue (2 or 10 kHz) and a response (pressing the left or right lever)25, setting the interval between the stimulus presentation and lever insertion to 3 s (Figure 1A). The success rate gradually increased, reaching a plateau around Day 13, which persisted through Day 24 (Figure 1B). Based on changes observed in the success rate, we categorized the learning processes into two phases: the acquisition phase, which lasted until Day 13; and the learned phase, which began after Day 13 and then continued at least until Day 24. In addition, we divided the acquisition phase into three stages depending on the success rate that included the early (< 60%), middle (60–80%), and late (80–100%) stages. We chose Days 2, 6, and 10 at the early, middle, and late stages in the acquisition phase, respectively, as well as Days 13 and 24 in the learned phase for statistical comparison of the success rate. There were significant differences between Days 2 and 6, Days 6 and 10, as well as Days 10 and 13, showing no significant difference between Days 13 and 24 (Figure 1C). In contrast, the response time and omission rate did not vary among the training days (Figures 1D and 1E), and the variation of response bias decreased along with the progress of auditory discrimination (Figure 1F).

Revised version: We employed a two-alternative auditory discrimination task, in which rats are required to make new associations between a tone instruction cue (high frequency tone of 10 kHz or low frequency tone of 2 kHz) and a response (pressing the right or left lever) (Nishizawa et al., 2012), setting the interval between the stimulus presentation and lever insertion to 3 s (Figure 1A). The success rate gradually increased, reaching a plateau around Day 13, which persisted through Day 24 (Figure 1B). Based on changes observed in the success rate, we categorized the learning processes into two phases: the acquisition phase, which lasted until Day 13; and the learned phase, which began after Day 13 and then continued at least until Day 24. In addition, we divided the acquisition phase into three stages depending on the success rate that included the early (< 60%), middle (60–80%), and late (80–100%) stages. We chose Days 2, 6, and 10 at the early, middle, and late stages in the acquisition phase, respectively, as well as Days 13 and 24 in the learned phase for statistical comparison of the success rate. There were significant differences between Days 2 and 6, Days 6 and 10, as well as Days 10 and 13, showing no significant difference between Days 13 and 24 (Figure 1C). In contrast, the response time and omission rate did not vary among the training days (Figure 1D and E), and the variation of response bias decreased along with the progress of auditory discrimination (Figure 1F).

Behavioral performance of two-alternative auditory discrimination learning.

(A) Schematic diagram of the auditory discrimination task. Each trial started by presentation of a tone instruction cue with the high (10 kHz) or low (2 kHz) frequency. Three seconds later, the room light was illuminated, and two retractable levers were inserted at the same time. The rats were required to press the right and left levers in response to the high and low tones, respectively. (B) Learning curve of the auditory discrimination in intact rats (n = 14). Inset indicates the cumulative curve of number achieved the success rate of more than 90%. (C) Success rate at Days 2, 6, 10, 13, and 24 (one-way repeated ANOVA, F[1.897,24.659] = 106.706, p = 1.1 × 10-12, post hoc Bonferroni test, Day 2 vs. Day 6, p = 0.001, Day 6 vs. Day 10, p = 3.6 × 10-4, Day 10 vs. Day 13, p = 0.040, and Day 13 vs. Day 24, p = 0.080). (D) Response time (one-way repeated ANOVA, F[2.203,28.642] = 2.386, p = 0.105). (E) Omission rate (F[4,52] = 0.699, p = 0.596). (F) Response bias (one-way repeated ANOVA, F[1.129,14.674] = 0.250, p = 0.653). Data are indicated as the mean ± s.e.m., and individual data are overlaid. *p < 0.05, **p < 0.01, and ***p < 0.001.

Distinct brain activity patterns in striatal subregions during learning processes

Earlier version: To investigate dynamic changes in regional brain activity in the entire striatum during auditory discrimination in the same animals, we conducted the positron emission tomography (PET) for small animals with 2-deoxy-2-[18F]fluoro-D-glucose (18F-FDG), which measures cerebral glucose metabolism correlated to brain activity26–28 (Figure 2A). Before the auditory discrimination task, we performed the single lever press task, in which a lever pseudorandomly presented on either the left or right side, and the rats learned to press a lever associated with receiving a reward. This task was used as a control that does not require the decision process based on the auditory stimulus. 18F-FDG-PET scanning was carried out on Day 4 of the single lever press task, and then during a series of training sessions including the acquisition (Days 2, 6, and 10) and learned (Day 24) phases (Figure 2B). The average number of lever presses in the single lever press task exceeded 80 times (81.6 ± 1.3, mean ± s.e.m.) on Day 2 and maintained that level until Day 4 (Figure 2C). In the following auditory discrimination task, the success rate gradually increased along with the training days (Figure 2D), showing a significant difference among sessions in which we performed the 18F-FDG-PET scanning (Figure 2E). The response time and omission rate were consistent among the days, and the variation of response bias became smaller as the task progresses (Figure 2E). These performances were similar to those in the intact rats (compared to Figures 1C–1F), suggesting that the procedure for 18F-FDG-PET scans does not affect the acquisition of discrimination.

Revised version: To investigate dynamic changes in regional brain activity in the entire striatum during auditory discrimination in the same animals, we conducted the positron emission tomography (PET) for small animals with 2-deoxy-2-[18F]fluoro-D-glucose (18F-FDG), which measures cerebral glucose metabolism correlated to brain activity (Cui et al., 2015; Kornblum et al., 2000; Zimmer et al., 2017) (Figure 2A). Before the auditory discrimination task, we performed the single lever press task, in which a lever pseudorandomly presented on either the left or right side, and the rats learned to press a lever associated with receiving a reward. This task was used as a control that does not require the decision process based on the auditory stimulus. 18F-FDG-PET scanning was carried out on Day 4 of the single lever press task, and then during a series of training sessions including the acquisition (Days 2, 6, and 10) and learned (Day 24) phases (Figure 2B). The average number of lever presses in the single lever press task exceeded 80 times (81.6 ± 1.3, mean ± s.e.m.) on Day 2 and maintained that level until Day 4 (Figure 2C). In the following auditory discrimination task, the success rate gradually increased along with the training days (Figure 2D), showing a significant difference among sessions in which we performed the 18F-FDG-PET scanning (Figure 2E). The response time and omission rate were consistent among the days, and the variation of response bias became smaller as the task progresses (Figure 2E). These performances were similar to those in the intact rats (compared to Figure 1C–F), suggesting that the procedure for 18F-FDG-PET scans does not affect the acquisition of discrimination.

Dissociable brain activity patterns during auditory discrimination learning among three striatal subregions.

Earlier version: (A) Schematic illustration of an awake rat received intravenously 18F-FDG injection through an indwelling catheter attached to the tail. The rat was conducted for the behavioral experiment and then used for microPET imaging. (B) Schedule for the behavioral training and 18F-FDG-PET scan. After 18F-FDG injection, rats (n = 14 rats) were subjected to the behavioral experiment (30 min) and returned to the home cage (15 min). Ten min before the PET scan started, the rats were anesthetized and then the scan was started (30 min). (C and D) Learning curves of the single lever press task (C) and auditory discrimination task (D). Arrowheads indicate the PET scan days. (E) Behavioral performance on the scan days during the auditory discrimination. Success rate (one-way repeated ANOVA, F[3,39] = 125.012, p = 4.8 × 10-20, post hoc Bonferroni test; Day 2 vs. Day 6, p = 5.0 × 10-7; Day 6 vs. Day 10, p = 0.001; Day 10 vs. Day 24, p = 0.003), response time (one-way repeated ANOVA, F[3,39] = 0.156, p = 0.926), omission rate (one-way repeated ANOVA, F[3,39] = 1.559, p = 0.215), and response bias (one-way repeated ANOVA, F[1.779,23.130] = 17.734, p = 3.6 × 10-5) are shown. (F) Representative images of coronal sections that compared the brain activity in the single lever press task with the activity on Day 2, 6, 10, or 24 in the discrimination task. Left panels show schematic illustrations of striatal subregions. (G) Horizontal (top) and coronal (bottom) images of striatal activation areas shown on Day 10 vs. single lever in (F). (H) Representative images of coronal section that compared the brain activity on Day 2 with that on Day 6, 10, or 24 in the discrimination task. Color bars indicate the T-values, and a value of 3.8 was used as the threshold corresponding to the uncorrected threshold (p < 0.001). (I) Schematic pictures showing voxel of interests for the aDLS, pVLS, and DMS (green, left hemisphere; and purple, right hemisphere). (J-L) Regional 18F-FDG uptakes in the aDLS (J, one-way repeated ANOVA, left aDLS, F[4,52] = 10.322, p = 3.0 × 10-6, post hoc Bonferroni test; single lever vs. Day 6, p = 0.016; Day 2 vs. Day 6, p = 6.5 × 10-5; Day 2 vs. Day 10, p = 0.019; Day 2 vs. Day 24, p = 0.017; Day 6 vs. Day 24, p = 0.041; right aDLS, F[4,52] = 7.462, p = 7.9 × 10-5, post hoc Bonferroni test, Day 2 vs. Day 6, p = 2.9 × 10-4; Day 2 vs. Day 10, p = 0.036; Day 6 vs. Day 24, p = 0.005), pVLS (K, one-way repeated ANOVA, left pVLS, F[2.368,30.784] = 4.152, p = 0.020, post hoc Bonferroni test; single lever vs. Day 10, p = 1.8 × 10-5; right pVLS, F[4,52] = 5.995, p = 4.8 × 10-4, post hoc Bonferroni test, single lever vs. Day 10, p = 0.001; Day 2 vs. Day 10, p = 0.011; Day 6 vs. Day 10, p = 0.032), and DMS (L, one-way repeated ANOVA, left DMS, F[4,52] = 12.836, p = 2.4 × 10-7, post hoc Bonferroni test; Day 2 vs. Day 24, p = 1.8 × 10-5; Day 6 vs. Day 24, p = 4.8 × 10-4; Day 10 vs. Day 24, p = 5.5 × 10-5; right DMS, F[4,52] = 10.717, p = 2.0 × 10-6, post hoc Bonferroni test; single lever vs. Day 24, p = 0.036; Day 2 vs. Day 24, p = 0.011; Day 6 vs. Day 24, p = 2.6 × 10-4; Day 10 vs. Day 24, p = 0.011). Arrowheads indicate the day with the most activations throughout the learning process. Data are indicated as the mean ± s.e.m., and individual data are overlaid. The anteroposterior coordinates from bregma (mm) are shown (F–I). Scale bar; 2 mm (I). *p < 0.05, **p < 0.01, and ***p < 0.001.

Revised version: (A) Schematic illustration of an awake rat received intravenously 18F- FDG injection through an indwelling catheter attached to the tail. The rat was conducted for the behavioral experiment and then used for microPET imaging. (B) Schedule for the behavioral training and 18F-FDG-PET scan. After 18F-FDG injection, rats (n = 14 rats) were subjected to the behavioral experiment (30 min) and returned to the home cage (15 min). Ten min before the PET scan started, the rats were anesthetized and then the scan was started (30 min). (C and D) Learning curves of the single lever press task (C) and auditory discrimination task (D). Arrowheads indicate the PET scan days. (E) Behavioral performance on the scan days during the auditory discrimination. Success rate (one-way repeated ANOVA, F[3,39] = 125.012, p = 4.8 × 10-20, post hoc Bonferroni test; Day 2 vs. Day 6, p = 5.0 × 10-7; Day 6 vs. Day 10, p = 0.001; Day 10 vs. Day 24, p = 0.003), response time (one-way repeated ANOVA, F[3,39] = 0.156, p = 0.926), omission rate (one-way repeated ANOVA, F[3,39] = 1.559, p = 0.215), and response bias (one-way repeated ANOVA, F[1.779,23.130] = 17.734, p = 3.6 × 10-5) are shown. (F) Representative images of coronal sections that compared the brain activity in the single lever press task with the activity on Day 2, 6, 10, or 24 in the discrimination task. Left panels show schematic illustrations of striatal subregions. (G) Horizontal (top) and coronal (bottom) images of striatal activation areas shown on Day 10 vs. single lever in (F). (H) Representative images of coronal section that compared the brain activity on Day 2 with that on Day 6, 10, or 24 in the discrimination task. Color bars indicate the T-values, and a value of 3.8 was used as the threshold corresponding to the uncorrected threshold (p < 0.001). (I) Schematic pictures showing voxel of interests for the aDLS, pVLS, and DMS. (J-L) Regional 18F- FDG uptakes in the aDLS (J, one-way repeated ANOVA, left aDLS, F[4,52] = 10.322, p = 3.0 × 10- 6, post hoc Bonferroni test; single lever vs. Day 6, p = 0.016; Day 2 vs. Day 6, p = 6.5 × 10-5; Day 2 vs. Day 10, p = 0.019; Day 2 vs. Day 24, p = 0.017; Day 6 vs. Day 24, p = 0.041; right aDLS, F[4,52] = 7.462, p = 7.9 × 10-5, post hoc Bonferroni test, Day 2 vs. Day 6, p = 2.9 × 10-4; Day 2 vs. Day 10, p = 0.036; Day 6 vs. Day 24, p = 0.005), pVLS (K, one-way repeated ANOVA, left pVLS, F[2.368,30.784] = 4.152, p = 0.020, post hoc Bonferroni test; single lever vs. Day 10, p = 1.8 × 10- 5; right pVLS, F[4,52] = 5.995, p = 4.8 × 10-4, post hoc Bonferroni test, single lever vs. Day 10, p = 0.001; Day 2 vs. Day 10, p = 0.011; Day 6 vs. Day 10, p = 0.032), and DMS (L, one-way repeated ANOVA, left DMS, F[4,52] = 12.836, p = 2.4 × 10-7, post hoc Bonferroni test; Day 2 vs. Day 24, p = 1.8 × 10-5; Day 6 vs. Day 24, p = 4.8 × 10-4; Day 10 vs. Day 24, p = 5.5 × 10-5; right DMS, F[4,52] = 10.717, p = 2.0 × 10-6, post hoc Bonferroni test; single lever vs. Day 24, p = 0.036; Day 2 vs. Day 24, p = 0.011; Day 6 vs. Day 24, p = 2.6 × 10-4; Day 10 vs. Day 24, p = 0.011). Arrowheads indicate the day with the most activations throughout the learning process. Data are indicated as the mean ± s.e.m., and individual data are overlaid. The anteroposterior coordinates from bregma (mm) are shown (F–I). Scale bar; 2 mm (I). *p < 0.05, **p < 0.01, and ***p < 0.001.

Earlier version: We investigated the brain regions, in which 18F-FDG uptake significantly increased or decreased during the learning processes, applying a statistical significance threshold (p < 0.001, uncorrected) and an extent threshold, by using a voxel-based statistical parametric analysis26. We first analyzed task-related brain activity by comparing 18F-FDG uptake on either day of the discrimination task with that on Day 4 of the single lever press task (summarized in Table S1). In the striatal subregions, there were significant increases of the activity in the unilateral aDLS and the bilateral pVLS on Days 6 and 10, respectively (Figures 2F and 2G). However, no significant change was observed in the DMS during the acquisition phase except a unilateral decrease on Day 24 (Figure 2F). Next, to evaluate the activity related to the progress of learning, 18F-FDG uptake on Days 6, 10, and 24 was compared with that on Day 2. The learning-dependent brain activity was significantly increased in the bilateral aDLS and the unilateral pVLS on Days 6 and 10, respectively (Figure 2H), whereas the activity decreased in the bilateral DMS on Day 24 (Figure 2H). This decrease was also detected when 18F-FDG uptakes were compared between Days 10 and 24 (Figure 2H), suggesting the down-regulation of the DMS activity in the learned phase.

Revised version: We investigated the brain regions, in which 18F-FDG uptake significantly increased or decreased during the learning processes, applying a statistical significance threshold (p < 0.001, uncorrected) and an extent threshold, by using a voxel-based statistical parametric analysis (Cui et al., 2015). We first analyzed task-related brain activity by comparing 18F-FDG uptake on either day of the discrimination task with that on Day 4 of the single lever press task (summarized in Figure 2–table supplement 1). In the striatal subregions, there were significant increases of the activity in the unilateral aDLS and the bilateral pVLS on Days 6 and 10, respectively (Figure 2F and G).

Added in the revised version: However, no significant change was observed in the DMS during the acquisition phase except a unilateral decrease on Day 24 (Figure 2F). Next, to evaluate the activity related to the progress of learning, 18F-FDG uptake on Days 6, 10, and 24 was compared with that on Day 2. The learning- dependent brain activity was significantly increased in the bilateral aDLS and the unilateral pVLS on Days 6 and 10, respectively (Figure 2H), whereas the activity decreased in the bilateral DMS on Day 24 (Figure 2H). This decrease was also detected when 18F-FDG uptakes were compared between Days 10 and 24 (Figure 2H), suggesting the down-regulation of the DMS activity in the learned phase.

To quantitatively validate changes in the brain activity in striatal subregions, we analyzed the amount of 18F-FDG uptake in the voxel of interest of these subregions (Figure 2I). The 18F- FDG uptake in the bilateral aDLS reached a peak on Day 6, and it decreased during the subsequent days (Figure 2J), whereas the uptake in the bilateral pVLS gradually increased along with the progress of learning, showing the maximal value on Day 10 (Figure 2K). The uptake in the DMS did not alter during the acquisition phase, and it decreased on Day 24 (Figure 2L). These results suggest that the aDLS and pVLS contribute to learning processes in a different temporal pattern during the acquisition phase, whereas the DMS may be engaged to the behavior in the learned phase.

Excitotoxic lesion of the aDLS and pVLS, but not the DMS, disrupts the acquisition of auditory discrimination

Earlier version: To study whether the changed activity of striatal subregions is involved in learning processes, rats received a bilateral injection of solution containing ibotenic acid (IBO; 8 mg/mL, 0.3 μL/site) or phosphate buffered saline (PBS) into the subregions, and we conducted for the behavioral study. The range of IBO lesion in the subregions was analyzed by immunostaining with a neuronal marker NeuN after the behavioral tests (Figures 3A–3C). The number of single lever presses increased similarly along with the training in both the PBS- and IBO-injected groups into the aDLS, pVLS, or DMS along with the training (Figures 3D–3F). For the aDLS injections, increases in success rate were impaired during the acquisition phase in the IBO group, compared to the PBS group, and the impairments were mainly observed at the middle stage (Figures 3G). For the pVLS injection, success rate impairments were observed in the IBO group showed compared to the PBS group, although the period showing a significant difference was unclear (Figure 3H). The response time in the aDLS-lesioned group did not alter during the discrimination training, whereas that in the pVLS-lesioned group was continuously prolonged (Figure. S1A and S1B), and the omission rate and response bias were similar between the lesioned and the corresponding control groups (Figures S1D/S1G for aDLS lesion and Figures S1E/S1H for pVLS lesion). In contrast, for the DMS injection the success rate in the discrimination task was comparable between the PBS and IBO groups (Figure 3I). The DMS lesion lengthened the response time through the learning (Figure S1C), whereas it did not alter the omission rate and response bias (Figures S1F and S1I). These results indicate that the aDLS and pVLS are necessary for the acquisition of auditory discrimination, and that the aDLS appears to be mainly involved in the learning process at the middle stage. The results also reveal that the DMS contributes to the execution, but not learning, of discriminative behavior.

Revised version: To study whether the changed activity of striatal subregions is involved in learning processes, rats received a bilateral injection of solution containing ibotenic acid (IBO; 8 mg/mL, 0.3 μL/site) or phosphate buffered saline (PBS) into the subregions, and we conducted for the behavioral study. The range of IBO lesion in the subregions was analyzed by immunostaining with a neuronal marker NeuN after the behavioral tests (Figure 3A–C). The number of single lever presses increased similarly along with the training in both the PBS- and IBO-injected groups into the aDLS, pVLS, or DMS along with the training (Figure 3D–F). For the aDLS injections, increases in success rate were impaired during the acquisition phase in the IBO group, compared to the PBS group, and the impairments were mainly observed at the middle stage (Figure 3G). For the pVLS injection, success rate impairments were observed in the IBO group showed compared to the PBS group, although the period showing a significant difference was unclear (Figure 3H). The response time in the aDLS-lesioned group did not alter during the discrimination training, whereas that in the pVLS- lesioned group was continuously prolonged (Figure 3–figure supplement 1A and B), and the omission rate and response bias were similar between the lesioned and the corresponding control groups (Figure 3–figure supplement 1D and G for aDLS lesion, and Figure 3 –figure supplement 1E and H for pVLS lesion). In contrast, for the DMS injection the success rate in the discrimination task was comparable between the PBS and IBO groups (Figure 3I). The DMS lesion lengthened the response time through the learning (Figure 3–figure supplement 1C), whereas it did not alter the omission rate and response bias (Figure 3–figure supplement 1F and I). These results indicate that the aDLS and pVLS are necessary for the acquisition of auditory discrimination, and that the aDLS appears to be mainly involved in the learning process at the middle stage. The results also reveal that the DMS contributes to the execution, but not learning, of discriminative behavior.

Impacts of excitotoxic lesion of striatal subregions on the acquisition of auditory discrimination.

Rats were given intracranial injection of PBS or IBO solution into the aDLS (n = 8 for each injection), pVLS (n = 8 for each injection), or DMS (n = 7 for PBS injection, and n = 6 for IBO injection) before start of the single lever press task. (A-C) Representative images of NeuN immunostaining and individual schematic illustrations showing lesioned area in the aDLS (A), pVLS (B), or DMS (C). Dotted lines in the images indicate the range of the lesioned area. ac, anterior commissure; cc, corpus callosum; and LV, lateral ventricle. (D-F) Learning curves during the single lever press task in the groups injected into the aDLS (D, two-way repeated ANOVA, group, F[1,14] = 1.275, p = 0.278, day, F[1.512,21.171] = 26.281, p = 7.0 × 10-6, group × day, F[1.512,21.171] = 0.834, p = 0.418), pVLS (E, two-way repeated ANOVA, group, F[1,14] = 2.542, p = 0.133, day, F[1.547,21.658] = 40.183, p = 2.2 × 10-7, group × day, F[1.547,21.658] = 0.953, p = 0.380), or DMS (F, two-way repeated ANOVA, group, F[1,11] = 0.025, p = 0.876, day, F[1.205,13.252] = 14.818, p = 0.001, group × day, F[1.205,13.252] = 0.610, p = 0.478). (G-I) Learning curves during the auditory discrimination task in the groups injected into the aDLS (G, two-way repeated ANOVA, group, F[1,14] = 11.578, p = 0.004, day, F[2.805,39.274] = 60.733, p = 1.8 × 10-14, group × day, F[2.805,39.274] = 3.863, p = 0.018, simple main effect, PBS vs. IBO, Day 1, p = 0.140, Day 2, p = 0.085, Day 3, p = 0.284, Day 4, p = 0.068, Day 5, p = 0.011, Day 6, p = 9.5 × 10-5, Day 7, p = 0.001, Day 8, p = 0.002, Day 9, p = 0.012, Day 10, p = 0.059, Day 11, p = 0.034, Day 12, p = 0.046, Day 13, p = 0.136, Day 14, p = 0.096), pVLS (H, two-way repeated ANOVA, group, F[1,14] = 7.221, p = 0.018, day, F[3.646,51.041] = 51.794, p = 9.2 × 10-17, group × day, F[3.646,51.041] = 1.612, p = 0.190), or DMS (I, two-way repeated ANOVA, group, F[1,11] = 0.469, p = 0.507, day, F[2.677,29.452] = 48.860, p = 3.9 × 10-11, group × day, F[2.677,29.452] = 1.094, p = 0.362). (J-L) Changes in the proportions of the win-shift-win or lose-shift-lose strategies through the discrimination learning in the groups treated into the aDLS (J, two-way repeated ANOVA, group, F[1,14] = 13.451, p = 0.003, day, F[3.048,42.673] = 54.777, p = 9.1 × 10-15, group × day, F[3.048,42.673] = 3.288, p = 0.029, simple main effect, PBS vs. IBO, Day 1, p = 0.728, Day 2, p = 0.225, Day 3, p = 0.165, Day 4, p = 0.044, Day 5, p = 0.015, Day 6, p = 1.1 × 10-4, Day 7, p = 0.002, Day 8, p = 5.1 × 10-4, Day 9, p = 0.004, Day 10, p = 0.060, Day 11, p = 0.053, Day 12, p = 0.081, Day 13, p = 0.198, Day 14, p = 0.075 for win-shift-win; group, F[1,14] = 5.039, p = 0.041, day, F[3.866,54.119] = 14.578, p = 5.2 × 10-8, group × day, F[3.866,54.119] = 2.471, p = 0.057 for lose-shift-lose), pVLS (K, two-way repeated ANOVA, group, F[1,14] = 9.251, p = 0.009, day, F[3.206,44.881] = 47.168, p = 2.7 × 10-14, group × day, F[3.206,44.881] = 1.862, p = 0.146 for win- shift-win; group, F[1,14] = 0.762, p = 0.397, day, F[4.773,66.815] = 24.018, p = 1.9 × 10-13, group × day, F[4.773,66.815] = 1.338, p = 0.260 for lose-shift-lose), or DMS (L, two-way repeated ANOVA, group, F[1,11] = 0.703, p = 0.420, day, F[3.052,33.568] = 48.424, p = 2.3 × 10-12, group × day, F[3.052,33.568] = 0.846, p = 0.480 for win-shift-win; group, F[1,11] = 0.002, p = 0.965, day, F[3.442,37.866] = 13.787, p = 1.0 × 10-6, group × day, F[3.442, 37.866] = 0.581, p = 0.654 for lose-shift-lose). Data are indicated as the mean ± s.e.m., and individual data are overlaid (except for panels J–L). The anteroposterior coordinates from bregma (mm) are shown (A–C). Scale bars; 2 mm (A–C). *p < 0.05, **p < 0.01 and ***p < 0.001. Abbreviations: WSW, win-shift-win; and LSL, lose-shift-lose.

Earlier version: Previous studies have reported that the dorsal striatum is involved in both behavioral strategies based on the stimulus-response association and the response-outcome association and during instrumental learning12,13. To assess the impact of striatal lesions on the behavioral strategies, we analyzed the proportions of responses attributed to two strategies in all responses of each session. One is the “win-shift-win” strategy, which is considered to reflect the behavioral strategy based on the stimulus-response association29,30. In this strategy, after a correct response in the previous trial, the rats press the opposite lever in the current trial in response to a shift of the instruction cue, resulting in the correct response. Another strategy is the “lose-shift-lose” strategy, which is considered to appear as a consequence of the behavioral strategy based on the response-outcome association29. In this strategy, after an error response in the previous trial, the rats press the opposite lever in the current trial, despite a shift of the instruction cue, leading to another error response. In the control group into either of three striatal subregions, the proportions of win-shift-win and lose-shift-lose strategies increased and decreased, respectively, during the acquisition phase (Figures 3J–3L). The changes in the proportion of the two strategies were similar to those in the intact rats (Figure S2). For the aDLS injection, the increases in the win-shift-win proportion were markedly impaired in the IBO group compared to the PBS group with significant differences mainly at the middle stage, and the decrease in the lose-shift-lose proportion was delayed compared to the same group (Figure 3J). For the pVLS-injection, the increasing win-shift-win proportion was also impaired in the IBO group compared to the PBS group, although the change in the lose-shift-lose proportion was comparable between the two groups (Figure 3K). In contrast, for the DMS lesion both win-shift-win and lose-shift-lose proportions in the IBO group did not show significant differences from the corresponding values in their control group (Figure 3L). These results indicate that the aDLS promotes the win-shift-win strategy through stimulus-response association and suppresses the lose-shift-lose strategy rooted in response-outcome association, and that the pVLS selectively facilitates the win-shift-win strategy acquisition. Thus, these two subregions not only play roles in the acquisition of auditory discrimination but also control distinct behavioral strategies during the learning processes.

Revised version: Previous studies have reported that the dorsal striatum is involved in both behavioral strategies based on the stimulus-response association and the response-outcome association and during instrumental learning (Dickinson, 1985; O’Hare et al., 2016). To assess the impact of striatal lesions on the behavioral strategies, we analyzed the proportions of responses attributed to two strategies in all responses of each session. One is the “win-shift-win” strategy, which is considered to reflect the behavioral strategy based on the stimulus-response association (Bergstrom et al., 2018; Packard et al., 1989). In this strategy, after a correct response in the previous trial, the rats press the opposite lever in the current trial in response to a shift of the instruction cue, resulting in the correct response. Another strategy is the “lose-shift-lose” strategy, which is considered to appear as a consequence of the behavioral strategy based on the response-outcome association (Bergstrom et al., 2018). In this strategy, after an error response in the previous trial, the rats press the opposite lever in the current trial, despite a shift of the instruction cue, leading to another error response. In the control group into either of three striatal subregions, the proportions of win-shift- win and lose-shift-lose strategies increased and decreased, respectively, during the acquisition phase (Figure 3J–L). The changes in the proportion of the two strategies were similar to those in the intact rats (Figure 1–figure supplement 1). For the aDLS injection, the increases in the win-shift- win proportion were markedly impaired in the IBO group compared to the PBS group with significant differences mainly at the middle stage, and the decrease in the lose-shift-lose proportion was delayed compared to the same group (Figure 3J). For the pVLS-injection, the increasing win- shift-win proportion was also impaired in the IBO group compared to the PBS group, although the change in the lose-shift-lose proportion was comparable between the two groups (Figure 3K). In contrast, for the DMS lesion both win-shift-win and lose-shift-lose proportions in the IBO group did not show significant differences from the corresponding values in their control group (Figure 3L). These results indicate that the aDLS promotes the win-shift-win strategy through stimulus- response association and suppresses the lose-shift-lose strategy rooted in response-outcome association, and that the pVLS selectively facilitates the win-shift-win strategy acquisition. Thus, these two subregions not only play roles in the acquisition of auditory discrimination but also control distinct behavioral strategies during the learning processes.

Earlier version: After the completion of discrimination training, we tested locomotion and feeding behavior of rats with aDLS and pVLS lesions. There was no difference in locomotor activity in the open field and the amount of food intake in both lesioned groups (Figures S1J-S1M), suggesting that learning deficits following excitotoxic lesions of the striatal subregions cannot be attributable to abnormalities in these general motor behaviors.

Revised version: After the completion of discrimination training, we tested feeding behavior and locomotion of rats with aDLS and pVLS lesions. There was no difference in the amount of food intake and locomotor activity in the open field in both lesioned groups (Figure 3–figure supplement 2A-D), suggesting that learning deficits following excitotoxic lesions of the striatal subregions cannot be attributable to abnormalities in these general motor behaviors.

Inhibiting aDLS and pVLS activity impairs discriminative behavior at different stages of the acquisition phase

Earlier version: Our 18F-FDG-PET imaging study indicated that the brain activity in the aDLS and pVLS was increased in a different temporal pattern during the acquisition phase of auditory discrimination. To examine whether neuronal activity in striatal subregions is linked with the processes at different stages, we temporarily inhibited the activity in each subregion by the treatment with a GABAA receptor agonist muscimol (MUS). In a pilot study, we tested the effect of the bilateral injection of several doses of MUS (0.1, 0.2, and 0.5 μg/μL) into the aDLS or pVLS on the single lever press behavior and confirmed that the administration of the lowest dose of MUS did not affect the frequency of single lever presses and the quantity of food intake (Figures S3A–S3D). In addition, to achieve the stage-specific inactivation of the aDLS or pVLS neurons, rats were given bilateral injections of MUS or saline (SAL) at three different timings during a series of the training. The first injection was conducted on Day 2 to target the early stage. The second and third injections were conducted on the day after the success rate had reached 60% and 80% for the first time through the training, respectively, to target the middle or late stage. After the injections, the rats were used for the behavioral test. Placement locations of the injection cannula into the striatal subregions were assessed by cresyl violet staining after the behavioral tests (Figures 4A and 4B). The success rates on the day before the injection (day N-1) at the early, middle, and late stages were similar between the two groups that received the bilateral injections of MUS or SAL (Figures S3E and S3F). Rats were then administered injections into the aDLS on the next day (Day N), and the success rate in the MUS-injected groups significantly decreased at the middle stage compared to the corresponding SAL-injected groups, but not at the early and late stages (Figure 4C, left). In contrast, the success rate in the MUS groups into the pVLS significantly decreased at the late stage, but not at the early and middle stages, compared with the respective SAL groups (Figure 4D, left). To normalize the variation of learning speed in the individual animals, we calculated the value by subtracting the success rate on Day N-1 from that on Day N in the individual animals. There were significant reductions in the subtracted rate at the middle stage in the MUS-injected group into the aDLS (Figure 4C, right) and at the late stage in the MUS-injected group into the pVLS (Figure 4D, right). These results indicate that the aDLS and pVLS mainly function at the middle and late stages, respectively, in the acquisition phase of auditory discrimination, supporting the results obtained from our imaging study.

Revised version: Our 18F-FDG-PET imaging study indicated that the brain activity in the aDLS and pVLS was increased in a different temporal pattern during the acquisition phase of auditory discrimination. To examine whether neuronal activity in striatal subregions is linked with the processes at different stages, we temporarily inhibited the activity in each subregion by the treatment with a GABAA receptor agonist muscimol (MUS). In a pilot study, we tested the effect of the bilateral injection of several doses of MUS (0.1, 0.2, and 0.5 μg/μL) into the aDLS or pVLS on the single lever press behavior and confirmed that the administration of the lowest dose of MUS did not change the frequency of single lever presses as compared to the saline (SAL) injection (Figure 4–figure supplement 1A and B). The lowest dose of MUS into the two striatal subregions unaffected the amount of food intake (Figure 4–figure supplement 1C and D). In addition, to achieve the stage- specific inactivation of the aDLS or pVLS neurons, rats were given bilateral injections of MUS (0.1 μg/μL, 0.25 μL per site) or SAL at three different timings during a series of the training. The first injection was conducted on Day 2 to target the early stage. The second and third injections were conducted on the day after the success rate had reached 60% and 80% for the first time through the training, respectively, to target the middle or late stage. After the injections, the rats were used for the behavioral test. Placement locations of the injection cannula into the striatal subregions were assessed by cresyl violet staining after the behavioral tests (Figure 4A and B). The success rates on the day before the injection (Day N-1) at the early, middle, and late stages were similar between the two groups that received the bilateral injections of MUS or SAL (Figure 4–figure supplement 2A and B). Rats were then administered injections into the aDLS on the next day (Day N), and the success rate in the MUS-injected groups significantly decreased at the middle stage compared to the corresponding SAL-injected groups, but not at the early and late stages (Figure 4C, left). In contrast, the success rate in the MUS groups into the pVLS significantly decreased at the late stage, but not at the early and middle stages, compared with the respective SAL groups (Figure 4D, left). To normalize the variation of learning speed in the individual animals, we calculated the value by subtracting the success rate on Day N-1 from that on Day N in the individual animals. There were significant reductions in the subtracted rate at the middle stage in the MUS-injected group into the aDLS (Figure 4C, right) and at the late stage in the MUS-injected group into the pVLS (Figure 4D, right). These results indicate that the aDLS and pVLS mainly function at the middle and late stages, respectively, in the acquisition phase of auditory discrimination, supporting the results obtained from our imaging study.

Influence of transient inhibition of striatal subregions at different timings on the performance of auditory discrimination.

Earlier version: Rats received intracranial injection of SAL or MUS solution into the aDLS (n = 8 for each injection) or pVLS (n = 6 for each injection). (A and B) Representative images of cresyl violet staining and schematic illustrations showing placement sites of the tip of guide cannula in the aDLS (A) or pVLS (B). Dotted lines in the images indicate the position of the cannula placement. Ac, anterior commissure; cc, corpus callosum; and LV, lateral ventricle. (C and B) Effects of transient striatal inhibition on the performance. For aDLS inhibition (C), success rate (left, two-way repeated ANOVA, stage, F[2,28] = 28.708, p = 1.7 × 10-7, group, F[1,14] = 8.840, p = 0.010, stage × group, F[2,28] = 0.757, p = 0.478; unpaired Student’s t-test, early, t[14] = 0.558, p = 0.586; middle, t[14] = 2.764, p = 0.015; Welch’s t-test, late, t[12.010] = 1.617, p = 0.132, p = 0.016 after a Bonferroni correction method) and subtracted success rate on day N-1 from that on day N (right, two-way repeated ANOVA, stage, F[2,28] = 1.231, p = 0.307, group, F[1,14] = 10.797, p = 0.005, stage × group, F[2,28] = 2.360, p = 0.113; unpaired Student’s t-test, early, t[14] = 0.285, p = 0.780, late, t[14] = 1.271, p = 0.225, Welch’s t-test, middle, t[9.603] = 3.331, p = 0.008, p = 0.016 after a Bonferroni correction method). For pVLS inhibition (D), success rate (left, two-way repeated ANOVA, stage, F[2,20] = 17.642, p = 3.8 × 10-5, group, F[1,10] = 43.942, p = 5.9 × 10-5, stage × group, F[2,20] = 4.729, p = 0.012; simple main effect, early, p = 0.055, middle, p = 0.109, late, p = 4.2 × 10-5) and subtracted success rate on day N-1 from that on day N (right, two-way repeated ANOVA, stage, F[2,20] = 12.105, p = 3.6 × 10-4, group, F[1,10] = 16.310, p = 0.002, stage × group, F[2,20] = 8.500, p = 0.002; simple main effect, early, p = 0.595, middle, p = 0.085, late, p = 3.0 × 10-6). (E and F) Subtracted proportion of the win-shift-win or lose-shift-lose strategies in the groups injected into the aDLS at the middle stage (E, unpaired Student’s t-test, t[14] = 2.038, p = 0.061 for win-shift-win, and t[14] = 2.714, p = 0.017 for lose-shift-lose) and the late stage (F, unpaired Student’s t-test, t[14] = 0.898, p = 0.384 for win-shift-win, and t[14] = 0.226, p = 0.824 for lose-shift-lose). (G and H) Subtracted proportion of the win-shift-win or lose-shift-lose strategies in the groups injected into the pVLS at the middle stage (G, unpaired Student’s t-test, t(10) = 1.924, p = 0.083 for win-shift-win, and Welch’s t-test, t[6.095] = 1.364, p = 0.221 for lose-shift-lose) and the late stage (H, unpaired Student’s t-test, t[10] = 3.629, p = 0.005 for win-shift-win, and t[10] = 1.577, p = 0.146 for lose-shift-lose). Data are indicated as the mean ± s.e.m., and individual data are overlaid. The anteroposterior coordinates from bregma (mm) are shown (A and B). Scale bars; 2 mm (A and B). *p < 0.05, **p < 0.01, and ***p < 0.001. Abbreviations: WSW, win-shift-win; and LSL, lose-shift-lose.

Revised version: Rats received intracranial injection of SAL or MUS solution into the aDLS (n = 8 for each injection) or pVLS (n = 6 for each injection). (A and B) Representative images of cresyl violet staining and schematic illustrations showing placement sites of the tip of guide cannula in the aDLS (A) or pVLS (B). Dotted lines in the images indicate the position of the cannula placement. ac, anterior commissure; cc, corpus callosum; and LV, lateral ventricle. (C and D) Effects of transient striatal inhibition on the performance. For aDLS inhibition (C), success rate (left, two-way repeated ANOVA, stage, F[2,28] = 28.708, p = 1.7 × 10-7, group, F[1,14] = 8.840, p = 0.010, stage × group, F[2,28] = 0.757, p = 0.478; unpaired Student’s t-test, early, t[14] = 0.558, p = 0.586; middle, t[14] = 2.764, p = 0.015; Welch’s t-test, late, t[12.010] = 1.617, p = 0.132, p = 0.016 after a Bonferroni correction method) and subtracted success rate on day N-1 from that on day N (right, two-way repeated ANOVA, stage, F[2,28] = 1.231, p = 0.307, group, F[1,14] = 10.797, p = 0.005, stage × group, F[2,28] = 2.360, p = 0.113; unpaired Student’s t- test, early, t[14] = 0.285, p = 0.780, late, t[14] = 1.271, p = 0.225, Welch’s t-test, middle, t[9.603] = 3.331, p = 0.008, p = 0.016 after a Bonferroni correction method). For pVLS inhibition (D), success rate (left, two-way repeated ANOVA, stage, F[2,20] = 17.642, p = 3.8 × 10-5, group, F[1,10] = 43.942, p = 5.9 × 10-5, stage × group, F[2,20] = 4.729, p = 0.012; simple main effect, early, p = 0.055, middle, p = 0.109, late, p = 4.2 × 10-5) and subtracted success rate on day N-1 from that on day N (right, two-way repeated ANOVA, stage, F[2,20] = 12.105, p = 3.6 × 10-4, group, F[1,10] = 16.310, p = 0.002, stage × group, F[2,20] = 8.500, p = 0.002; simple main effect, early, p = 0.595, middle, p = 0.085, late, p = 3.0 × 10-6). (E and F) Subtracted proportion of the win-shift-win or lose- shift-lose strategies in the groups injected into the aDLS at the middle stage (E, unpaired Student’s t-test, t[14] = 2.038, p = 0.061 for win-shift-win, and t[14] = 2.714, p = 0.017 for lose-shift-lose) and the late stage (F, unpaired Student’s t-test, t[14] = 0.898, p = 0.384 for win-shift-win, and t[14] = 0.226, p = 0.824 for lose-shift-lose). (G and H) Subtracted proportion of the win-shift-win or lose- shift-lose strategies in the groups injected into the pVLS at the middle stage (G, unpaired Student’s t-test, t(10) = 1.924, p = 0.083 for win-shift-win, and Welch’s t-test, t[6.095] = 1.364, p = 0.221 for lose-shift-lose) and the late stage (H, unpaired Student’s t-test, t[10] = 3.629, p = 0.005 for win- shift-win, and t[10] = 1.577, p = 0.146 for lose-shift-lose). Data are indicated as the mean ± s.e.m., and individual data are overlaid. The anteroposterior coordinates from bregma (mm) are shown (A and B). Scale bars; 2 mm (A and B). *p < 0.05, **p < 0.01, and ***p < 0.001. Abbreviations: WSW, win-shift-win; and LSL, lose-shift-lose.

Earlier version: Next, we analyzed the influence of temporal striatal inactivation on the behavioral strategy at the middle and late stages of the acquisition phase. We calculated the proportions of the win-shift-win and lose-shift-lose strategies, and subtracted the values of these proportions on Day N-1 from those on Day N. The proportion of both strategies on Day N-1 were comparable between the MUS-and SAL-injected groups (Figures S3G and S3H). In the aDLS, the subtracted proportion of the lose-shift-lose strategy at the middle stage significantly increased in the MUS group as compared to the SAL group, although the subtracted proportion of the win-shift-win strategy tended to decrease in the MUS group as compared to the control group (Figure 4E), and the subtracted proportions of the win-shift-win or lose-shift-lose strategy at the late stage was similar between both groups (Figure 4F). In the pVLS, the subtracted proportions of the win-shift-win and lose-shift-lose strategies at the middle stage were similar between the MUS and SAL groups (Figure 4G), and the subtracted win-shift-win proportion at the late stage significantly decreased in the MUS group as compared to the SAL group, although the subtracted lose-shift-lose proportion was comparable between the two groups (Figure 4H). Our results show that transient inhibition of aDLS activity at the middle stage affects the lose-shift-lose strategy and has a subtle effect on the win-shift-win strategy, supporting that the aDLS mediates the acquisition process of auditory discrimination at the middle stage at least partly through the suppression of the behavior based on the response-outcome association. Although the impact of inhibition on the win-shift-win strategy was subtle, excitotoxic lesion of the aDLS influenced the win-shift-win strategy in addition to the lose-shift-lose strategy (see Figure 3J), suggesting that the aDLS acts to enhance the auditory discrimination by promoting the stimulus-response association, together with the suppression of the response-outcome association. The discrepancy in the extent of changes in the win-shift-win strategy between the two treatments may be generated from the difference in the timing or sustainability of drug inhibition of the striatal subregion. In addition, our results show that pVLS inactivation at the late stage selectively impairs the win-shift-win strategy, supporting contribution of the pVLS to the process via progression of the stimulus-response association.

Revised version: Next, we analyzed the influence of temporal striatal inactivation on the behavioral strategy at the middle and late stages of the acquisition phase. We calculated the proportions of the win- shift-win and lose-shift-lose strategies, and subtracted the values of these proportions on Day N-1 from those on Day N. The proportion of both strategies on Day N-1 were comparable between the MUS- and SAL-injected groups (Figure 4–figure supplement 3A and B). In the aDLS, the subtracted proportion of the lose-shift-lose strategy at the middle stage significantly increased in the MUS group as compared to the SAL group, although the subtracted proportion of the win-shift-win strategy tended to decrease in the MUS group as compared to the control group (Figure 4E), and the subtracted proportions of the win-shift-win or lose-shift-lose strategy at the late stage was similar between both groups (Figure 4F). In the pVLS, the subtracted proportions of the win-shift- win and lose-shift-lose strategies at the middle stage were similar between the MUS and SAL groups (Figure 4G), and the subtracted win-shift-win proportion at the late stage significantly decreased in the MUS group as compared to the SAL group, although the subtracted lose-shift-lose proportion was comparable between the two groups (Figure 4H). Our results show that transient inhibition of aDLS activity at the middle stage affects the lose-shift-lose strategy and has a subtle effect on the win-shift-win strategy, supporting that the aDLS mediates the acquisition process of auditory discrimination at the middle stage at least partly through the suppression of the behavior based on the response-outcome association. Although the impact of inhibition on the win-shift-win strategy was subtle, excitotoxic lesion of the aDLS influenced the win-shift-win strategy in addition to the lose-shift-lose strategy (see Figure 3J), suggesting that the aDLS acts to enhance the auditory discrimination by promoting the stimulus-response association, together with the suppression of the response-outcome association. The discrepancy in the extent of changes in the win-shift-win strategy between the two treatments may be generated from the difference in the timing or sustainability of drug inhibition of the striatal subregion. In addition, our results show that pVLS inactivation at the late stage selectively impairs the win-shift-win strategy, supporting contribution of the pVLS to the process via progression of the stimulus-response association.

Neuronal activity in the aDLS and pVLS during the acquisition phase

Earlier version: To examine the firing activity in striatal subregions during the acquisition phase of auditory discrimination, we simultaneously recorded neuronal activity in the aDLS and pVLS at the unilateral (left) side by using a multi-unit recording procedure (Figure 5A). Using immunostaining for tyrosine hydroxylase (TH) after the recordings, we verified the location of electrodes within the striatal subregions (Figure 5B). After the single lever press task, we conducted the auditory discrimination task for the recordings. In correct trials, the delay period (0.5 ± 0.2 s) was inserted between a lever press and a sound indicating the reward to distinguish neuronal activity related to the two events (Figure 5C). The behavioral data in individual rats were divided into three grouped sessions corresponding to the early, middle, and late stages, based on a sigmoid function analysis31 of the success rate in the task. The success rates in the grouped sessions were 52.7 ± 0.8 %, 61.6 ± 0.7 %, and 85.3 ± 1.7 % (mean ± s.e.m.) for the early, middle, and late stages, respectively, and the rate was significantly increased throughout the three stages (Figure S4A). However, the response time, omission rate, and response bias were similar among the stages (Figures S4B-S4D). Thus, these stages corresponded to the timing window defined in the acquisition phase of the transient MUS inhibition experiment of neuronal activity.

Revised version: To examine the firing activity in striatal subregions during the acquisition phase of auditory discrimination, we simultaneously recorded neuronal activity in the aDLS and pVLS at the unilateral (left) side by using a multi-unit recording procedure (Figure 5A). Using immunostaining for tyrosine hydroxylase (TH) after the recordings, we verified the location of electrodes within the striatal subregions (Figure 5B). After the single lever press task, we conducted the auditory discrimination task for the recordings. In correct trials, the delay period (0.5 ± 0.2 s) was inserted between a lever press and a sound indicating the reward to distinguish neuronal activity related to the two events (Figure 5C). The behavioral data in individual rats were divided into three grouped sessions corresponding to the early, middle, and late stages, based on a sigmoid function analysis (Bathellier et al., 2013) of the success rate in the task. The success rates in the grouped sessions were 52.7 ± 0.8 %, 61.6 ± 0.7 %, and 85.3 ± 1.7 % (mean ± s.e.m.) for the early, middle, and late stages, respectively, and the rate was significantly increased throughout the three stages (Figure 5– figure supplement 1A). However, the response time, omission rate, and response bias were similar among the stages (Figure 5–figure supplement 1B-D). Thus, these stages corresponded to the timing window defined in the acquisition phase of the transient MUS inhibition experiment of neuronal activity.

Multi-unit recording of neurons in striatal subregions and firing activity related to the behavioral outcome of reward sound-related HR type neurons in the aDLS.

Earlier version: (A) Schematic illustration of a freely moving rat used for simultaneous multi-unit recordings in the aDLS and pVLS. Six rats were used for the following analysis. (B) Representative images showing positions of electrode tips (arrows) in the aDLS and pVLS (left rows). The recording sites estimated by electrode tracks and electrical marks in individual rats are shown (center and right rows). ac, anterior commissure; LV, lateral ventricle. (C) Sequence of some events in correct and error trials. The delay periods were pseudorandomly added between a correct lever press and the reward sound (0.5 ± 0.2 s). The room light turned off the extended time (4 s) after the correct response or immediately after the error response. (D) Mean firing rate (top rows) and auROC values (bottom rows) of reward sound-related HR type neurons in the aDLS at the early (n = 43 neurons), middle (n = 37 neurons), and late (n = 44 neurons) stages. These data are aligned to the choice response. Green shadows show the period when the reward sound is presented (500 ± 200 ms after the lever press). Time bins with significant differences between the mean firing rate in the HR or LR trial and any of the rates in other three trials (Wilcoxon signed rank test) or the distribution of auROC values and 0.5 (Wilcoxon rank-sum test) are represented by the circles at the top. (E) Averaged firing rate during the reward sound period in the HR and LR trials at the early, middle, and late stages (Kruskal-Wallis test: HR trial, χ2 = 3.028, p = 0.220; and LR trial, χ2 = 1.249, p = 0.536). (F) Cumulative probability of the proportion of reward sound-related HR type neuron number at the three stages (Kruskal-Wallis test, χ2 = 45.378, p = 1.4 × 10-10; post hoc Tukey-Kramer test, p = 7.1 × 10-5 for early vs. middle, p = 1.0 × 10-9 for early vs. late, and p = 0.040 for middle vs. late). Data are indicated as the mean ± s.e.m. (D) or the median and quartiles with the maximal and minimal values (E). The anteroposterior coordinates from bregma (mm) are shown (B). Scale bars; 2 mm (B). *p < 0.05 and ***p < 0.001.

Revised version: (A) Schematic illustration of a freely moving rat used for simultaneous multi-unit recordings in the aDLS and pVLS. Six rats were used for the following analysis. (B) Representative images of TH immunostaining, showing positions of electrode tips (arrows) in the aDLS and pVLS (left rows). The recording sites estimated by electrode tracks and electrical marks in individual rats are shown (center and right rows). ac, anterior commissure; LV, lateral ventricle. (C) Sequence of some events in correct and error trials. The delay periods were pseudorandomly added between a correct lever press and the reward sound (0.5 ± 0.2 s). The room light turned off the extended time (4 s) after the correct response or immediately after the error response. (D) Mean firing rate (top rows) and auROC values (bottom rows) of reward sound-related HR type neurons in the aDLS at the early (n = 43 neurons), middle (n = 37 neurons), and late (n = 44 neurons) stages. These data are aligned to the choice response. Green shadows show the period when the reward sound is presented (500 ± 200 ms after the lever press). Time bins with significant differences between the mean firing rate in the HR or LR trial and any of the rates in other three trials (Wilcoxon signed rank test) or the distribution of auROC values and 0.5 (Wilcoxon rank-sum test) are represented by the circles at the top. (E) Averaged firing rate during the reward sound period in the HR and LR trials at the early, middle, and late stages (Kruskal-Wallis test: HR trial, χ2 = 3.028, p = 0.220; and LR trial, χ2 = 1.249, p = 0.536). (F) Cumulative probability of the proportion of reward sound-related HR type neuron number at the three stages (Kruskal-Wallis test, χ2 = 45.378, p = 1.4 × 10-10; post hoc Tukey- Kramer test, p = 7.1 × 10-5 for early vs. middle, p = 1.0 × 10-9 for early vs. late, and p = 0.040 for middle vs. late). Data are indicated as the mean ± s.e.m. (D) or the median and quartiles with the maximal and minimal values (E). The anteroposterior coordinates from bregma (mm) are shown (B). Scale bars; 2 mm (B). *p < 0.05 and ***p < 0.001.

Earlier version: We identified 1,062 and 423 well-isolated neurons from the aDLS and pVLS, respectively, across the three stages (Figure S4E) and focused on the major population of striatal neurons, putative medium spiny neurons (MSNs; Figures S4F and S4G). The total numbers of MSNs at the early, middle, and late stages were 295, 364, and 360 neurons for the aDLS, and 87, 110, and 190 neurons for pVLS, respectively. Based on the combinatorial pattern of the tone instruction cue and lever press in our discrimination task, we categorized the electrophysiological data into four trial types: (1) high-frequency tone/right lever press (HR) and (2) low-frequency tone/left lever press (LL) as correct responses; and (3) low-frequency tone/right lever press (LR) and (4) high-frequency tone/left lever press (HL) as error responses (see Figure 1A). We identified HR or LL type neurons showing significant changes in firing rate related to the cue onset, choice response, reward sound, or first licking as compared to the baseline firing rate in each of two trials (Mann-Whitney U test, p < 0.01). These neurons were further divided into two groups based on increased or decreased activity (Figures S5A and S5B). In the following analyses, we focused on the HR and LL type neurons with increased event-related activity to explore the firing patterns of neurons underlying behavioral processes of auditory discrimination.

Revised version: We identified 1,062 and 423 well-isolated neurons from the aDLS and pVLS, respectively, across the three stages (Figure 5–figure supplement 2A) and focused on the major population of striatal neurons, putative medium spiny neurons (MSNs; Figure 5–figure supplement 2B). The numbers of MSNs were 96% and 91% of the total numbers of identified neurons in the aDLS and pVLS, respectively (Figure 5-figure supplement 2C). The neuron numbers at the early, middle, and late stages were 295, 364, and 360 neurons in the aDLS, and 87, 110, and 190 neurons in the pVLS, respectively. Based on the combinatorial pattern of the tone instruction cue and lever press in our discrimination task, we categorized the electrophysiological data into four trial types: (1) high- frequency tone/right lever press (HR) and (2) low-frequency tone/left lever press (LL) as correct responses; and (3) high-frequency tone/left lever press (HL) and (4) low-frequency tone/right lever press (LR) as error responses (see Figure 1A). We identified HR or LL type neurons showing significant changes in firing rate related to the cue onset, choice response, reward sound, or first licking as compared to the baseline firing rate in each of two trials (Mann-Whitney U test, p < 0.01). These neurons were further divided into two groups based on increased or decreased activity (Figure 5–figure supplement 3A and B). In the following analyses, we focused on the HR and LL type neurons with increased event-related activity to explore the firing patterns of neurons underlying behavioral processes of auditory discrimination.

Subpopulations of aDLS neurons, but not of pVLS neurons, show firing activity related to the behavioral outcome

Earlier version: Previous studies have reported that neurons in the dorsal striatum show the outcome-related activity in the instrumental conditioning tasks24,32. We thus investigated whether aDLS and pVLS neurons exhibit a response related to the behavioral outcome during the acquisition of auditory discrimination. In the aDLS, the firing activity of reward sound-related HR type neurons showed a higher extent in the increase during a period, when the reward sound is presented in the correct responses, in the HR trial as compared with the LR, LL, or HL trial throughout the three stages (Figure 5D, top rows). The rate of these neurons also moderately increased during the same period in the LR trial compared to the LL or HL trial across the three stages, although the extent of the increase was lower than that in the HR trial (Figure 5D, top rows). The area under the receiver operating characteristic (auROC) values revealed a significant difference of the HR trial from the LR trial during the reward sound period, whereas the values of HL and LL trials indicated a significant difference in the opposite direction from the LR trial during the same period (Figure 5D, bottom rows). The averaged firing rate during the reward sound period in the HR or LR trial appeared to modestly increase at the middle stage, but no significant difference existed among the stages (Figure 5E). To directly compare the number of neurons showing event-related activity among different stages, we calculated the number of reward sound-HR neurons by matching the trial number at each stage (Figure 5F). The proportion decreased progressively from the early to middle stages, and to the late stage. In addition, the firing rate of reward sound-related LL type neurons exhibited a higher increase in the LL trial and a modest increase in the HL trial while the averaged firing rate during the reward sound period was similar among the three stages. Furthermore, the proportion of reward sound-LL neurons reduced progressively through the stages (Figures S6A–S6C). These results indicate that the aDLS contained neurons related to the outcome of choice, showing a greater increase in the firing activity for the correct responses (HR and LL trials) than the error responses (LR and HL trials), and that the proportion of these subpopulations was gradually reduced through the progress of learning. In contrast, during the reward sound period at all stages, the firing rate and auROC values of reward sound-HR and reward sound-LL neurons in the pVLS revealed no significant differences across the four kinds of trial (Figures S6D and S6E), indicating that the firing activity of pVLS neurons did not appear to be associated with the behavioral outcome.

Revised version: Previous studies have reported that neurons in the dorsal striatum show the outcome-related activity in the instrumental conditioning tasks (Guo et al., 2019; Nonomura et al., 2018). We thus investigated whether aDLS and pVLS neurons exhibit a response related to the behavioral outcome during the acquisition of auditory discrimination. In the aDLS, the firing activity of reward sound- related HR type neurons showed a higher extent in the increase during a period, when the reward sound is presented in the correct responses, in the HR trial as compared with the LL, HL, or LR trial throughout the three stages (Figure 5D, top rows). The rate of these neurons also moderately increased during the same period in the LR trial compared to the LL or HL trial across the three stages, although the extent of the increase was lower than that in the HR trial (Figure 5D, top rows). The area under the receiver operating characteristic (auROC) values revealed a significant difference of the HR trial from the LR trial during the reward sound period, whereas the values of LL and HL trials indicated a significant difference from the LR trial in the opposite direction during the same period (Figure 5D, bottom rows). The averaged firing rate during the reward sound period in the HR or LR trial appeared to modestly increase at the middle stage, but no significant difference existed among the stages (Figure 5E). To directly compare the number of neurons showing event-related activity among different stages, we calculated the number of reward sound- HR type neurons by matching the trial number at each stage (Figure 5F). The proportion decreased progressively from the early to middle stages, and to the late stage. In addition, the firing rate of reward sound-related LL type neurons exhibited a higher increase in the LL trial and a modest increase in the HL trial while the averaged firing rate during the reward sound period was similar among the three stages (Figure 5–figure supplement 4A and B). Furthermore, the proportion of reward sound-LL neurons reduced progressively through the stages (Figure 5–figure supplement 4C). These results indicate that the aDLS contained neurons related to the outcome of choice, showing a greater increase in the firing activity for the correct responses (HR and LL trials) than the error responses (HL and LR trials), and that the proportion of these subpopulations was gradually reduced through the progress of learning. In contrast, during the reward sound period at all stages, the firing rate and auROC values of reward sound-HR and reward sound-LL type neurons in the pVLS revealed no significant differences across the four kinds of trial (Figure 5–figure supplement 5A and B), indicating that the firing activity of pVLS neurons did not appear to be associated with the behavioral outcome.

Subpopulations of aDLS and pVLS neurons show sustained firing activity after the reward in different manners

Earlier version: The dorsal striatum is known to demonstrate a response of neuronal activity after the positive outcome in the instrumental behaviors33,34. We thus examined whether aDLS and pVLS neurons show the response after obtaining the reward during the acquisition phase of auditory discrimination. In the aDLS, the firing rate of the first licking-related HR type neurons indicated a sustained elevation after the first licking in the HR trial as compared with any of the other three trials across the early, middle, and late stages (Figure 6A, top rows). The auROC values revealed a significant difference of the LR, LL, or HL trial from the HR trial across the three stages (Figure 6A, middle rows). In addition, we calculated the averaged firing rate and increased firing period, defined as the total number of time bins above 3 s.d. of the baseline firing rate for 5 s after the first licking in the HR trial. The averaged firing rate was comparable between the early and middle stages, but the rate was significantly lower at the late stage compared to either of the other two stages (Figure 6B). The increased firing period was also comparable between the early and middle stages, showing a significant reduction at the late stage (Figure 6C). The proportions of first licking-HR neurons were similar between the early and late stages, although they showed a slight decrease at the middle stage (Figure 6D). There were few neurons with significant correlations between the firing rate and licking performance (Figure 6E). In addition, the firing rate of the first licking-related LL type neurons exhibited the largest increase in the LL trial at the three stages, showing a tendency to reduce the average firing rate and increased firing period after the first licking across the stages (Figures S7A-S7C). The proportion of first licking-LL neurons during the learning processes displayed similar changes as observed in the first licking-HR neurons (Figure S7D). These results indicate that the aDLS contained neurons showing long-lasting, increased activity after the reward in the specific correct combination of the stimulus and response (HR or LL trial), and that the extent of firing activity and length of firing period in the response appeared to be gradually reduced with no correlation between the proportion of neurons and the progress of the learning stages.

Revised version: The dorsal striatum is known to demonstrate a response of neuronal activity after the positive outcome in the instrumental behaviors (Smith and Graybiel, 2016; Williams and Eskandar, 2006). We thus examined whether aDLS and pVLS neurons show the response after obtaining the reward during the acquisition phase of auditory discrimination. In the aDLS, the firing rate of the first licking-related HR type neurons indicated a sustained elevation after the first licking in the HR trial as compared with any of the other three trials across the early, middle, and late stages (Figure 6A, top rows). The auROC values revealed a significant difference of the LL, HL, or LR trial from the HR trial across the three stages (Figure 6A, middle rows). In addition, we calculated the averaged firing rate and increased firing period, defined as the total number of time bins above 3 s.d. of the baseline firing rate for 5 s after the first licking in the HR trial. The averaged firing rate was comparable between the early and middle stages, but the rate was significantly lower at the late stage compared to either of the other two stages (Figure 6B). The increased firing period was also comparable between the early and middle stages, showing a significant reduction at the late stage (Figure 6C). The proportions of first licking-HR type neurons were similar between the early and late stages, although they showed a slight decrease at the middle stage (Figure 6D). There were only a few neurons with significant correlations between the firing rate and licking performance (Figure 6E). In addition, the firing rate of the first licking-related LL type neurons exhibited the largest increase in the LL trial at the three stages, showing a tendency to reduce the average firing rate and increased firing period after the first licking across the stages (Figure 6–figure supplement 1A-C). The proportion of first licking-LL type neurons during the learning processes displayed similar changes as observed in the first licking-HR neurons, and there were few neurons with significant correlations between the firing rate and licking performance (Figure 6–figure supplement 1D and E). These results indicate that the aDLS contained neurons showing long- lasting, increased activity after the reward in the specific correct combination of the stimulus and response (HR or LL trial), and that the extent of firing activity and length of firing period in the response appeared to be gradually reduced with no correlation between the proportion of neurons and the progress of the learning stages.

Sustained activity after the reward of first licking-related HR type neurons in the aDLS and pVLS.

Earlier version: (A) Firing rate (top rows), auROC (middle rows), and licking rate (bottom rows) after the reward of first licking-related HR type neurons in the aDLS at the early (n = 27 neurons), middle (n = 31 neurons), and late (n = 49 neurons) stages. These data are aligned to the choice response. Time bins with significant differences between the mean firing rate in the HR trial and any of the rate in other three trials (Wilcoxon signed rank test) or the distribution of auROC values and 0.5 (Wilcoxon rank-sum test) are represented by the circles at the top. Timings of the reward sound and first licking are shown. (B) Averaged firing rate for 5 s after the first licking in the HR trial (Kruskal-Wallis test, χ2 = 16.396, p = 2.8 × 10-4, post hoc Tukey-Kramer test, early vs. middle, p = 0.726, early vs. late, p = 7.2 × 10-4, middle vs. late, p = 0.009). (C) Averaged total number of time bins above 3 S.D. of the baseline firing of aDLS neurons for 5 s after the first licking in the HR trial (Kruskal-Wallis test, χ2 = 12.803, p = 0.002, post hoc Tukey-Kramer test, early vs. middle, p = 0.173, early vs. late, p = 0.001, middle vs. late, p = 0.219). (D) Cumulative probability of the proportion of first licking-related HR type neuron number (Kruskal-Wallis test, χ2 = 39.280, p = 3.0 × 10-9, post hoc Tukey-Kramer test, early vs. middle, p = 1.3 × 10-7, early vs. late, p = 0.995, middle vs. late, p = 2.3 × 10-7). (E) Distribution of the correlation coefficient between the numbers of spikes of first licking-related HR type neurons and licking. Closed column indicates the number of neurons showing significant correlations between the two parameters (5 out of 107 neurons). (F) Firing rate (top rows), auROC (middle rows), and licking rate (bottom rows) after the reward of FL-HR neurons in the pVLS at the early (n = 17 neurons), middle (n = 11 neurons), and late (n = 64 neurons) stages. These data are aligned to the choice response. Time bins with significant differences between the rate in the HR trial and any of the rate in other trials (Wilcoxon signed rank test) or the distribution of auROC values and 0.5 (Wilcoxon rank-sum test) are represented by the circles at the top. The timings of the reward sound and first licking are shown. (G) Averaged firing rate of pVLS neurons after the first licking in the HR trial (Kruskal-Wallis test, χ2 = 7.652, p = 0.022; post hoc Tukey-Kramer test, early vs. middle, p = 0.975, early vs. late, p = 0.076, and middle vs. late, p = 0.096). (H) Averaged total number of time bins above 3 S.D. of the baseline firing after the first licking in the HR trial (Kruskal-Wallis test, χ2 = 0.023, p = 0.989). (I) Cumulative probability of the proportion of first licking-related HR type neuron number (Kruskal-Wallis test, χ2 = 47.922, p = 3.9 × 10-11; post hoc Tukey-Kramer test, early vs. middle, p = 0.011, early vs. late, p = 1.7 × 10-4, and middle vs. late, p = 9.7 × 10-10). (J) Distribution of the correlation coefficient between the number of spikes in the pVLS and the licking. Closed column indicates the number of neurons showing significant correlations between the two parameters (5 out of 92 neurons). Data are indicated as the mean ± s.e.m. (A and F) or the median and quartiles with the maximal and minimal values (B, C, G, and H). **p < 0.01 and ***p < 0.001. Abbreviations: RS, reward sound; and FL, first licking.

Revised version: (A) Firing rate (top rows), auROC (middle rows), and licking rate (bottom rows) after the reward of first licking-related HR type neurons in the aDLS at the early (n = 27 neurons), middle (n = 31 neurons), and late (n = 49 neurons) stages. These data are aligned to the choice response. Timings of the reward sound and first licking are shown. Time bins with significant differences between the mean firing rate in the HR trial and any of the rate in other three trials (Wilcoxon signed rank test) or the distribution of auROC values and 0.5 (Wilcoxon rank-sum test) are represented by the circles at the top. (B) Averaged firing rate for 5 s after the first licking in the HR trial (Kruskal-Wallis test, χ2 = 16.396, p = 2.8 × 10-4, post hoc Tukey-Kramer test, early vs. middle, p = 0.726, early vs. late, p = 7.2 × 10-4, middle vs. late, p = 0.009). (C) Averaged total number of time bins above 3 s.d. of the baseline firing for 5 s after the first licking in the HR trial (Kruskal-Wallis test, χ2 = 12.803, p = 0.002, post hoc Tukey-Kramer test, early vs. middle, p = 0.173, early vs. late, p = 0.001, middle vs. late, p = 0.219). (D) Cumulative probability of the proportion of first licking-related HR type neuron number (Kruskal-Wallis test, χ2 = 39.280, p = 3.0 × 10-9, post hoc Tukey-Kramer test, early vs. middle, p = 1.3 × 10-7, early vs. late, p = 0.995, middle vs. late, p = 2.3 × 10-7). (E) Distribution of the correlation coefficient between the numbers of spikes and licking. Closed column indicates the number of neurons showing significant correlations between the two parameters (5 out of 107 neurons). (F) Firing rate (top rows), auROC (middle rows), and licking rate (bottom rows) after the reward of first licking-related HR neurons in the pVLS at the early (n = 17 neurons), middle (n = 11 neurons), and late (n = 64 neurons) stages. The data are aligned to the choice response, and the timings of the reward sound and first licking are shown. Time bins with significant differences between the rate in the HR trial and any of the rate in other trials or the distribution of auROC values and 0.5 are represented as above. (G) Averaged firing rate of pVLS neurons after the first licking in the HR trial (Kruskal-Wallis test, χ2 = 7.652, p = 0.022; post hoc Tukey-Kramer test, early vs. middle, p = 0.975, early vs. late, p = 0.076, and middle vs. late, p = 0.096). (H) Averaged total number of time bins above 3 s.d. of the baseline firing of pVLS neurons after the first licking in the HR trial (Kruskal-Wallis test, χ2 = 0.023, p = 0.989). (I) Cumulative probability of the proportion of first licking-related HR type neuron number (Kruskal-Wallis test, χ2 = 47.922, p = 3.9 × 10-11; post hoc Tukey-Kramer test, early vs. middle, p = 0.011, early vs. late, p = 1.7 × 10-4, and middle vs. late, p = 9.7 × 10-10). (J) Distribution of the correlation coefficient between the number of spikes and the licking. Closed column indicates the number of neurons showing significant correlations between the two parameters (5 out of 92 neurons). Data are indicated as the mean ± s.e.m. (A and F) or the median and quartiles with the maximal and minimal values (B, C, G, and H). **p < 0.01 and ***p < 0.001. Abbreviations: RS, reward sound; and FL, first licking.

Earlier version: In the pVLS, the first licking-HR neurons showed a sustained elevation in the firing rate after receiving a reward in the HR trial compared to that in the LR or HL trial throughout the three stages (Figure 6F, top rows). A similar elevation in the firing rate was also observed in the LL trial throughout the stages (Figure 6F, top rows). The auROC values revealed a significant difference of the LR or HL trial, but not the LL trial, from the HR trial throughout the stages (Figure 6F, middle rows). The average firing rate did not show significant changes among the three stages (Figure 6G), and the increased firing period was also similar across the stages (Figure 6H). The proportion of first licking-HR neurons increased at the late stage compared to the early or middle stage (Figure 6I). The firing rate showed almost no significant correlation to the licking performance (Figure 6J). In addition, the first licking-LL neurons exhibited similar firing patterns and changes in the proportion of neurons across the stages (Figures S7F–S7J). Therefore, the pVLS included neurons representing the long-lasting increased activity after the reward in the two kinds of correct responses (HR and LL trials) with the increased number of these neurons at the late stage of learning.

Revised version: In the pVLS, the first licking-HR type neurons showed a sustained elevation in the firing rate after receiving a reward in the HR trial compared to that in the HL or LR trial throughout the three stages (Figure 6F, top rows). A similar elevation in the firing rate was also observed in the LL trial throughout the stages (Figure 6F, top rows). The auROC values revealed a significant difference of the HL or LR trial, but not the LL trial, from the HR trial throughout the stages (Figure 6F, middle rows). The average firing rate did not show significant changes among the three stages (Figure 6G), and the increased firing period was also similar across the stages (Figure 6H).

Added in the revised version: The proportion of first licking-HR neurons increased at the late stage compared to the early or middle stage (Figure 6I). The firing rate showed almost no significant correlation to the licking performance (Figure 6J). In addition, the first licking-LL neurons exhibited similar firing patterns and changes in the proportion of neurons across the stages (Figure 6–figure supplement 2A-E).

Added in the revised version: Therefore, the pVLS included neurons representing the long-lasting increased activity after the reward in the two kinds of correct responses (HR and LL trials) with the increased number of these neurons at the late stage of learning.

Neurons in the pVLS, but not in the aDLS, exhibit firing patterns related to the beginning and ending of the behavior

Earlier version: The dorsal striatum plays a key role in the acquisition and execution of behaviors based on the stimulus-response association, and the association is accompanied by a transient activation of neuronal firing immediately after the stimulus presentation or action in trials, which is considered to represent the beginning and ending of a learned behavior35–37. We first examined whether the firing activity of cue onset-related HR type or cue onset-related LL type neurons corresponds to the beginning of the behavior in the auditory discrimination task. In the aDLS, the firing rate of the cue onset-HR and cue onset-LL neurons showed a significant difference at some time points in each trial compared to those in any of the other three trials, with a significant difference of auROC values between the trials (Figures S8A and S8C). However, the transiently increased activity did not emerge at the time points corresponding to the beginning of the behavior. In contrast, the cue onset-HR neurons in the pVLS showed a transient increase in the firing rate immediately after the cue onset, and the activity appeared at the middle and late stages (Figure 7A, top). The transient activation was observed not only in the HR trial but also in the other three trials, being consistent with the results of a previous report38. The increased firing rate showed a slight but notable difference between the HR and LL trials with a significant difference in the auROC values between the two trials (Figure 7A, low), suggesting that different combinations between the stimulus and response may affect the level of firing activity. Differences in the averaged firing rate between the two kinds of trial was observed at the middle and late stages (Figure 7B). The proportion of cue onset-HR neurons increased at the middle and late stages compared to the early stage (Figure 7C). In addition, the firing rate of the cue onset-LL neurons exhibited patterns similar to those observed in the cue onset-HR neurons (Figure S8E). The data indicate that subpopulations of pVLS neurons show firing activity related to the beginning of the behavior, with a rising number of neurons throughout discrimination learning.

Revised version: The dorsal striatum plays a key role in the acquisition and execution of behaviors based on the stimulus-response association, and the association is accompanied by a transient activation of neuronal firing immediately after the stimulus presentation or action in trials, which is considered to represent the beginning and ending of a learned behavior (Barnes et al., 2005; Jin and Costa, 2010; Jog et al., 1999). We first examined whether the firing activity of cue onset-related HR and LL type neurons corresponds to the beginning of the behavior in the auditory discrimination task. In the aDLS, the firing rate of both types of neurons showed a significant difference at some time points in each trial compared to those in any of the other three trials, with a significant difference of auROC values between the trials (Figure 7–figure supplement 1A and B). However, the transiently increased activity did not emerge at the time points corresponding to the beginning of the behavior. In contrast, the cue onset-related HR type neurons in the pVLS showed a transient increase in the firing rate immediately after the cue onset, and the activity appeared at the middle and late stages (Figure 7A, top rows). The transient activation was observed not only in the HR trial but also in the other three trials, being consistent with the results of a previous report (Barnes et al., 2011). The increased firing rate showed a slight but notable difference between the HR and LL trials with a significant difference in the auROC values between the two kinds of trials (Figure 7A, bottom rows), suggesting that different combinations between the stimulus and response may affect the level of firing activity. The rate differences between the two trials were observed at the middle and late stages (Figure 7B). The proportion of cue onset-HR neurons increased at the middle and late stages compared to the early stage (Figure 7C). In addition, the firing rate of the cue onset-related LL type neurons exhibited patterns similar to those observed in the cue onset-HR neurons (Figure 7–figure supplement 2A). The data indicate that subpopulations of pVLS neurons show firing activity related to the beginning of the behavior, with a rising number of neurons throughout discrimination learning.

Earlier version: Transient activity related to the beginning and ending of a behavior of cue onset-HR and choice response-related HR type neurons in the pVLS.

Revised version: Transient activity related to the beginning and ending of a behavior of cue onset- and choice response-related HR type neurons in the pVLS.

Earlier version: (A) Firing rate (top rows) and auROC values (bottom rows) of cue onset-related HR type neurons in the pVLS at the early (n = 4 neurons), middle (n = 21 neurons), and late (n = 24 neurons) stages. These data are aligned to the cue onset. Time bins with significant differences between the mean firing rate in the HR trial and either of the rate in other three trials (Wilcoxon signed rank test) or the distribution of auROC values and 0.5 (Wilcoxon rank-sum test) are represented by the circles at the top. (B) Averaged firing rate of the cue onset-related HR type neurons in HR and LL trials (0-600 ms after the cue onset) (Wilcoxon signed rank test; early, p = 0.875, middle, p = 0.006, and late, p = 0.004). (C) Cumulative probability of the proportion of the cue onset-related HR type neuron number (Kruskal-Wallis test, χ2 = 51.114, p = 8.0 × 10-12; post hoc Tukey-Kramer test, early vs. middle, p = 9.6 × 10- 10, early vs. late, p = 0.002, and middle vs. late, p = 6.4 × 10-4). (D) Firing rate (top rows) and auROC values (bottom rows) of the choice response-related HR type neurons in the pVLS at the early (n = 11 neurons), middle (n = 22 neurons), and late (n = 41 neurons) stages. These data are aligned to the choice response. Time bins with significant differences between the mean rate in the HR trial and any of the rate in other trials (Wilcoxon signed rank test) or the distribution of auROC values and 0.5 (Wilcoxon rank-sum test) are represented by the circles at the top. The timing of the reward sound is presented. (E) Averaged firing rate of the choice response-related HR type neurons in HR and LL trials (0-600 ms after the choice response) (Wilcoxon signed rank test; early, p = 0.278, middle, p = 0.758, and late, p = 0.279). (F) Cumulative probability of the proportion of the choice response-related HR type neuron number (Kruskal-Wallis test, χ2 = 13.786, p = 0.001; post hoc Tukey-Kramer test, early vs. middle, p = 0.003, early vs. late, p = 0.005, and middle vs. late, p = 0.976). Data are indicated as the mean ± s.e.m. (A and B) or the median and quartiles with the maximal and minimal values (B and E). **p < 0.01 and ***p < 0.001. Abbreviation: RS, reward sound.

Revised version: (A) Firing rate (top rows) and auROC values (bottom rows) of cue onset-related HR type neurons in the pVLS at the early (n = 4 neurons), middle (n = 21 neurons), and late (n = 24 neurons) stages. These data are aligned to the cue onset. Time bins with significant differences between the mean firing rate in the HR trial and either of the rate in other three trials (Wilcoxon signed rank test) or the distribution of auROC values and 0.5 (Wilcoxon rank-sum test) are represented by the circles at the top. (B) Averaged firing rate of the cue onset-related HR type neurons in HR and LL trials (0-600 ms after the cue onset) (Wilcoxon signed rank test; early, p = 0.875, middle, p = 0.006, and late, p = 0.004). (C) Cumulative probability of the proportion of the cue onset-related HR type neuron number (Kruskal- Wallis test, χ2 = 51.114, p = 8.0 × 10-12; post hoc Tukey-Kramer test, early vs. middle, p = 9.6 × 10- 10, early vs. late, p = 0.002, and middle vs. late, p = 6.4 × 10-4). (D) Firing rate (top rows) and auROC values (bottom rows) of the choice response-related HR type neurons in the pVLS at the early (n = 11 neurons), middle (n = 22 neurons), and late (n = 41 neurons) stages. These data are aligned to the choice response. The timing of the reward sound is presented. Time bins with significant differences between the mean rate in the HR trial and any of the rate in other trials or the distribution of auROC values and 0.5 are represented as above. (E) Averaged firing rate of the choice response-related HR type neurons in HR and LL trials (0-600 ms after the choice response) (Wilcoxon signed rank test; early, p = 0.278, middle, p = 0.758, and late, p = 0.279). (F) Cumulative probability of the proportion of the choice response-related HR type neuron number (Kruskal- Wallis test, χ2 = 13.786, p = 0.001; post hoc Tukey-Kramer test, early vs. middle, p = 0.003, early vs. late, p = 0.005, and middle vs. late, p = 0.976). Data are indicated as the mean ± s.e.m. (A and D) or the median and quartiles with the maximal and minimal values (B and E). **p < 0.01 and ***p < 0.001. Abbreviation: RS, reward sound.

Earlier version: Next, we examined the firing activity of choice response-related HR type or choice response-related LL type neurons, corresponding to the ending of a behavior in auditory discrimination. In the aDLS, the firing rate of these neurons did not exhibit the transient increase related to the ending of the behavior, instead, there was a modest increase before the lever press (Figures S8B and S8D). In contrast, the choice response-HR neurons in the pVLS were transiently activated immediately after the lever press in the four kinds of trial at the middle and late stages, and the rate appeared to gradually decline toward the reward sound presentation (Figure 7D). The averaged firing rate was similar between the HR and LL trials though the stages (Figure 7E). The proportion of the choice response-HR neurons increased at the middle and late stages compared to the early stage (Figure 7F). In addition, the rate of choice response-LL neurons displayed similar firing patterns to those observed in the choice response-HR neurons (Figure S8F). Thus, subpopulations of pVLS neurons also represent the activity related to the ending of the behavior in auditory discrimination, increasing the neuron number throughout the learning stages in a similar pattern to neurons relevant to the beginning of the behavior in the task.

Revised version: Next, we examined the firing activity of choice response-related HR and LL type neurons, corresponding to the ending of a behavior in auditory discrimination. In the aDLS, the firing rate of these neurons did not exhibit the transient increase related to the ending of the behavior, instead, there was a modest increase before the lever press (Figure 7–figure supplement 3A and B). In contrast, the choice response-related HR type neurons in the pVLS were transiently activated immediately after the lever press in the four kinds of trial at the middle and late stages, and the rate appeared to gradually decline toward the reward sound presentation (Figure 7D). The averaged firing rate was similar between the HR and LL trials though the stages (Figure 7E). The proportion of the choice response-HR neurons increased at the middle and late stages compared to the early stage (Figure 7F). In addition, the rate of choice response-related LL type neurons displayed similar firing patterns to those observed in the choice response-HR neurons (Figure 7–figure supplement 2B). Thus, subpopulations of pVLS neurons also represent the activity related to the ending of the behavior in auditory discrimination, increasing the neuron number throughout the learning stages in a similar pattern to neurons relevant to the beginning of the behavior in the task.

Added in the revised version: Structure of neural circuits linked to the aDLS and pVLS

Added in the revised version: Anatomically, striatal subregions are known to form heterogeneous cortico-basal ganglia loops in mice (Foster et al., 2021; Hintiryan et al., 2016; Hunnicutt et al., 2016). To characterize the structure of neural circuits connected to the aDLS and pVLS in rats, a solution containing cholera toxin subunit B (CTb) labeled with Alexa 488 or Alexa 555 (1.0 mg/mL, 0.3 μL/site) as a bidirectional tracer (Amita et al., 2019) was injected into the aDLS and pVLS of the same animals, respectively (Figure 8A). Brain sections were prepared, and cell bodies and axon terminals were detected by CTb with florescence signals. Sections through the aDLS and pVLS indicated the florescence signals around the injection sites (Figure 8B). Alexa 488 signals were localized in cell bodies in the bilateral, rostrodorsal regions of some cortical areas, including the primary motor cortex, secondary motor cortex, primary somatosensory cortex, and secondary somatosensory cortex, and in the ipsilateral regions of various thalamic nuclei, such as the centrolateral nucleus, centromedian nucleus, parafascicular nucleus, and ventromedial nucleus, whereas Alexa 555 signals were observed in the bilateral, ventrolateral regions of some cortical areas, including the most rostral parts of the primary motor cortex, agranular/dysgranular/granular insular cortices, rhinal cortex, and amygdala nucleus, as well as in the ipsilateral regions of various thalamic nuclei, such as the centromedian nucleus, parafascicular nucleus, and medial geniculate nucleus (Figure 8C and D). In addition, Alexa 488 signals were distributed in axon terminals in the anterior part of the globus pallidus (GP) and entopeduncular nucleus (EPN), whereas Alexa 555 signals were detected in the posterior GP (Figure 8E). In the ventral midbrain. Alexa 488 and 555 signals in the nerve terminals were distributed distinctively in the ventromedial and dorsolateral parts of the substantia nigra pars reticulata (SNr), respectively, along with the rostrocaudal axis (Figure 8F). In the cell bodies, both signals were localized in the centromedial and lateral parts of the substantia nigra pars compacta (SNc), respectively (Figure 8F). The signals in the SNc were colocalized with immunoreactivities for TH, a marker of dopamine neurons in the ventral midbrain (Figure 8G).

Added in the revised version: Segregation of neural circuits linked to the aDLS and pVLS.

Added in the revised version: (A) Schematic illustration of the major inputs to and outputs from the striatum. A solution containing CTb labeled with Alexa 488 (green) or Alexa 555 (magenta) was injected into the aDLS and pVLS of the same animals (n = 2). (B) Coordinates for tracer injection into the aDLS and pVLS. Right, representative fluorescence images of sections through the injection sites. (C) Typical images of sections through the cerebral cortex and intralaminar thalamus. (D) Distribution of cell bodies innervating the aDLS and pVLS. Fluorescence signals in the cell bodies were marked by green or magenta dots using ImageJ (version 1.53) and overlaid on gray-scaled images of a brain section. (E) Axonal terminals arising from the aDLS and pVLS to the GP or EPN. Representative images of tracing are shown in the left panels. Camera lucida drawings of labeled signals are shown in the right panels. (F) Labeling of axonal terminals and cell bodies in the ventral midbrain. Representative images are shown in the left panels. Upper images are magnified views of the rectangles in the lower image. Camera lucida drawings are shown in the right panels. Arrowheads in the magnified images indicate labeled cell bodies. (G) Triple-fluorescence histochemistry with sections through the ventral midbrain. The sections were stained with an anti-TH antibody and then with a secondary IgG conjugated with Alexa 647. Representative images of the centromedial (left) and lateral (right) parts of the SNc are shown. Arrowheads indicate TH-positive cells with CTb-labeled signals in each image. The anteroposterior coordinates from bregma (mm) are shown. Scale bars; 1 mm (B-D), 500 μm (E), 100 μm (F and G). Abbreviations: AI, agranular insular cortex; Amy, amygdala nucleus; Au, auditory cortex; CL, centrolateral nucleus; CM, centromedian nucleus; DI, dysgranular insular cortex; GI, granular insular cortex; M1, primary motor cortex; M2, secondary motor cortex; MG, medial geniculate nucleus; PF, parafascicular nucleus; Rhi, rhinal cortex; S1, primary somatosensory cortex; S2, secondary somatosensory cortex; TeA, temporal association cortex; VM, ventromedial nucleus; and VTA, ventral tegmental area.

Added in the revised version: These results indicate that the aDLS and pVLS receive innervations from different subregions in the cerebral cortex, intralaminar thalamus, and SNc, and send projections to different subregions in the basal ganglia nuclei including the GP, EPN, and SNr, being consistent with previous anatomical studies of cortico-basal ganglia loops (Foster et al., 2021; Hintiryan et al., 2016; Hunnicutt et al., 2016).

Discussion

Earlier version: Our findings demonstrate that the aDLS and pVLS are required for the acquisition of auditory discrimination at different stages of learning, and that the DMS is not necessary for the acquisition of discrimination, although it appears to contribute to the execution of actions, in particular to the regulation of response time. These results not only support some recent studies reporting no requirement of the DMS in the procedural learning20,39,40, but also reveal important roles of two different subregions in the lateral striatum in the external cue-dependent decision-making. We propose that aDLS and pVLS neurons act to integrate the new learning of auditory discrimination in spatiotemporally and functionally distinct manners, showing a remarkable difference from the prior learning model that proposes the functional dominance changes from the DMS to DLS subregions during procedural learning12–16.

Revised version: Our findings demonstrate that the aDLS and pVLS are required for the acquisition of auditory discrimination at different stages of learning, and that the DMS is not necessary for the acquisition of discrimination, although it appears to contribute to the execution of actions, in particular to the regulation of response time. These results not only support some recent studies reporting no requirement of the DMS in the procedural learning (Featherstone and McDonald, 2004; Smith et al., 2021; Wolff et al., 2022), but also reveal important roles of two different subregions in the lateral striatum in the external cue-dependent decision-making. We propose that aDLS and pVLS neurons act to integrate the new learning of auditory discrimination in spatiotemporally and functionally distinct manners, showing a remarkable difference from the prior learning model that proposes the functional dominance changes from the DMS to DLS subregions during procedural learning (Dickinson, 1985; O’Hare et al., 2016; Redgrave et al., 2010; Yin et al., 2004: Yin et al., 2005).

Earlier version: Our 18F-FDG-PET imaging study demonstrated that the brain activity in the aDLS reaches a peak at the middle stage of the acquisition phase of auditory discrimination (Figure 2I). The results of excitotoxic lesion experiment of striatal subregions suggest that the aDLS functions to enhance the acquisition of discrimination through the promotion of behavior based on the stimulus-response association and the suppression of behavior based on the response-outcome association (Figure 3J). Transient inhibition experiment of neuronal activity confirmed this aDLS function at the middle stage of the learning, although the drug inhibition tended to reduce the stimulus-response behavior (Figures 4E and 4F). In addition, we found subpopulations of aDLS neurons (reward sound-related neurons) which represent the behavioral outcome and differentiate the correct (HR or LL) and incorrect (LR or HL) choices (Figures 5D and S6A). The distinct responses of these neurons may explain the mechanism by which the aDLS regulates the stimulus-response and response-outcome associations through distinct neural circuits leading to a differential control of the two associations. Alternatively, the responses mainly promote the behavior based on the stimulus-response association, thereby secondarily inhibiting the response-outcome association. Indeed, some previous studies have suggested that an action based on either the stimulus-response association or response-outcome association is generated as a result of competition between different striatal subregions41. We also found other subpopulations of aDLS neurons (first licking-related neurons) showing long-lasting increased activity after obtaining reward in response to a specific combination between the stimulus and response (HR or LL trial) (Figures 6A and S7A). This activity may function as the feedback signal to promote the stimulus and response association, resulting in the enhancement of discrimination learning. The proportion of neurons related to the behavioral outcome showed a gradual reduction through the progress of auditory discrimination (Figures 5E and S6C), and the extent of the firing activity and the length of the firing period of neurons showing sustained activation after reward were also decreased at the late stage of learning (Figures 6B and 6C). The learning-dependent changes in the properties of these subpopulations in the aDLS may explain the mechanism underlying the earlier acquisition process of the auditory discrimination.

Revised version: Our 18F-FDG-PET imaging study demonstrated that the brain activity in the aDLS reaches a peak at the middle stage of the acquisition phase of auditory discrimination (Figure 2J). The results of excitotoxic lesion experiment of striatal subregions suggest that the aDLS functions to enhance the acquisition of discrimination through the promotion of behavior based on the stimulus- response association and the suppression of behavior based on the response-outcome association (Figure 3G and J). Transient inhibition experiment of neuronal activity confirmed this aDLS function at the middle stage of the learning, although the drug inhibition tended to reduce the stimulus-response behavior (Figure 4C, E and F). In addition, we found subpopulations of aDLS neurons (reward sound-related neurons) which represent the behavioral outcome and differentiate the correct (HR or LL) and incorrect (HL or LR) choices (Figure 5D and Figure 5–figure supplement 4A). The distinct responses of these neurons may explain the mechanism by which the aDLS regulates the stimulus-response and response-outcome associations through distinct neural circuits leading to a differential control of the two associations. Alternatively, the responses mainly promote the behavior based on the stimulus-response association, thereby secondarily inhibiting the response-outcome association. Indeed, some previous studies have suggested that an action based on either the stimulus-response association or response-outcome association is generated as a result of competition between different striatal subregions (Gremel et al., 2016). We also found other subpopulations of aDLS neurons (first licking-related neurons) showing long-lasting increased activity after obtaining reward in response to a specific combination between the stimulus and response (HR or LL trial) (Figures 6A and Figure 6–figure supplement 1A). This activity may function as the feedback signal to promote the stimulus and response association, resulting in the enhancement of discrimination learning. The proportion of neurons related to the behavioral outcome showed a gradual reduction through the progress of auditory discrimination (Figures 5F and Figure 5–figure supplement 4C, and the extent of the firing activity and the length of the firing period of neurons showing sustained activation after reward were also decreased at the late stage of learning (Figure 6B and C). The learning-dependent changes in the properties of these subpopulations in the aDLS may explain the mechanism underlying the earlier acquisition process of the auditory discrimination.

Earlier version: The 18F-FDG-PET imaging indicated that the pVLS activity was gradually elevated through the acquisition phase of auditory discrimination, showing the maximal value at the late stage (Figure 2J). The lesion experiment suggested that the pVLS drives the formation of discrimination learning via the selective progression of the stimulus-response association (Figure 3K), and the transient inhibition experiment ascertained the pVLS function at the late stage of learning (Figures 4G and 4H). Previous studies have shown that the stimulus-response association is encoded in the DLS or VLS as the transient neuronal activity representing the beginning and ending of a learned behavior35–37,42. A recent study reported that the VLS is causally related to learned performance based on the auditory stimulus, while the DLS shows a weaker relationship to the behavior22. In the present study, we found subpopulations of pVLS neurons that represent the beginning and ending of a behavior (cue onset- and choice response-related neurons, respectively) (Figures 7A, 7D, S8E, and S8F), showing an increased number of neurons along with the progress of discrimination learning (Figures 7C and 7F). These data suggest that pVLS neurons contribute to the formation of the stimulus-response association through the learning processes. In addition, we found another subpopulation of pVLS neurons (first licking-related neurons) that are continuously activated after the reward outcome in response to two kinds of correct choices (HR and LL trials) (Figures 6F and S7F), and the proportion of neurons was increased at the late stage (Figures 6I and S7I). This activity may act as a feedback signal to consolidate the stimulus and response association, leading to the maintenance of associative memory at the late stage during the acquisition phase of learning.

Revised version: The 18F-FDG-PET imaging indicated that the pVLS activity was gradually elevated through the acquisition phase of auditory discrimination, showing the maximal value at the late stage (Figure 2K). The lesion experiment suggested that the pVLS drives the formation of discrimination learning via the selective progression of the stimulus-response association (Figure 3H and K), and the transient inhibition experiment ascertained the pVLS function at the late stage of learning (Figure 4D, G and H). Previous studies have shown that the stimulus-response association is encoded in the DLS or VLS as the transient neuronal activity representing the beginning and ending of a learned behavior (Barnes et al., 2005; Chen et al., 2021; Jin and Costa, 2010; Jog et al., 1999). A recent study reported that the VLS is causally related to learned performance based on the auditory stimulus, while the DLS shows a weaker relationship to the behavior (Lee et al., 2020). In the present study, we found subpopulations of pVLS neurons that represent the beginning and ending of a behavior (cue onset- and choice response-related neurons, respectively) (Figure 7A and D, Figure 7–figure supplement 2A and B), showing an increased number of neurons along with the progress of discrimination learning (Figure 7C and F). These data suggest that pVLS neurons contribute to the formation of the stimulus-response association through the learning processes. In addition, we found another subpopulation of pVLS neurons (first licking-related neurons) that are continuously activated after the reward outcome in response to two kinds of correct choices (HR and LL trials) (Figure 6F and Figure 6–figure supplement 2A), and the proportion of neurons was increased at the late stage (Figure 6I and Figure 6–figure supplement 2D). This activity may act as a feedback signal to consolidate the stimulus and response association, leading to the maintenance of associative memory at the late stage during the acquisition phase of learning.

Earlier version: In the present study, we identified neurons with four kinds of event-related activity in the aDLS and pVLS, and classified them into the HR and LL type neurons in which their firing activity showed a significant increase in the HR and LL trials, respectively, as compared to the baseline (Figures 5–7 and S6-S8). Parts of these two types of neurons with the respective event-related activity overlapped, and the numbers of overlapping neurons varied between striatal subregions (Figures S9A and S9B). In the pVLS, the overlapping neurons for the same event-related activities were 35.1%, 45.2%, and 67.0% in the cue onset, choice response, and first licking, respectively. In contrast, the aDLS had lower overlaps of 12.6% and 22.7% in the reward sound and first licking, respectively. Thus, distinct neurons in the aDLS tend to have properties of a single trial type while the same neurons in the pVLS seem to possess properties of both types. One possible explanation for the diversity in the proportion of overlapping neurons between the two subregions may be the difference of input patterns into striatal neurons or plastic changes in functional connectivity during learning processes. The detailed mechanism by which aDLS and pVLS neurons react to the events in a different manner should be investigated in the future.

Revised version: In the present study, we identified neurons with four kinds of event-related activity in the aDLS and pVLS, and classified them into the HR and LL type neurons (Figure 5-7, Figure 5- figure supplement 4 and 5, Figure 6-figure supplement 1 and 2, and Figure 7-figure supplement 1-3). Parts of these two types of neurons with the respective event-related activity overlapped, and the numbers of overlapping neurons varied between striatal subregions. In the pVLS, the overlapping neurons for the same event-related activities were 35.1%, 45.2%, and 67.0% in the cue onset, choice response, and first licking, respectively. In contrast, the aDLS had lower overlaps of 12.6% and 22.7% in the reward sound and first licking, respectively. Thus, distinct neurons in the aDLS tend to have properties of a single trial type while the same neurons in the pVLS seem to possess properties of both types. One possible explanation for the diversity in the proportion of overlapping neurons between the two subregions may be the difference of input patterns into striatal neurons or plastic changes in functional connectivity during learning processes. The detailed mechanism by which aDLS and pVLS neurons react to the events in a different manner should be investigated in the future.

Earlier version: Anatomically, striatal subregions are known to form heterogeneous cortico-basal ganglia loops in mice2,43,44. We examined the neural circuits linked to the aDLS and pVLS in rats by using bidirectional tracers. The results showed that two striatal subregions receive innervations from different subregions in the cerebral cortex, intralaminar thalamus, and substantia nigra pars compacta, and send projections to different subregions in the basal ganglia nuclei, including the globus pallidus, entopeduncular nucleus, and substantia nigra pars reticulata (Figure S10), confirming that the aDLS and pVLS form the heterogenous cortico-basal ganglia loops in the rat brain. These data suggest that the aDLS and pVLS mediate the acquisition of auditory discrimination via the parallel loops, and cooperate to achieve neural processing related to the dual regulation of the stimulus-response and response-outcome associations. Further studies are needed to elucidate the mechanism by which the two kinds of loops are processed separately and integrate the new learning of discrimination.

Revised version: We examined the neural circuits linked to the aDLS and pVLS in rats by using bidirectional tracers. The results showed that these two striatal subregions receive innervations from different subregions in the cerebral cortex, intralaminar thalamus, and ventral midbrain, and provide projections to different subregions in the basal ganglia nuclei (Figure 8), confirming that the two subregions form the heterogenous cortico-basal ganglia loops in the rodent brain (Foster et al., 2021; Hintiryan et al., 2016; Hunnicutt et al., 2016). These data suggest that the aDLS and pVLS mediate the acquisition of auditory discrimination via the parallel loops, and cooperate to achieve neural processing related to the dual regulation of the stimulus-response and response-outcome associations. Further studies are needed to elucidate the mechanism by which the two kinds of loops are processed separately and integrate the new learning of discrimination.

Limitations

Added in the revised version: In the 18F-FDG PET neuroimaging study, the brain activity in the aDLS reached a peak at the middle stage of the acquisition phase of auditory discrimination (Figure 2J). In the multi-unit electrophysiological recording experiment, the firing activity of aDLS neuron subpopulations related to the behavioral outcome showed no significant differences among the three stages (Figure 5E and Figure 5-figure supplement 4B), while the proportions of these subpopulations gradually reduced through the progress of learning stages (Figure 5F and Figure 5-figure supplement 4C). The extent of firing activity and length of firing period of other subpopulations showing sustained activation after the reward appeared to show learning-dependent decrease (Figure 6B and C, and Figure 6-figure supplement 1B and C), although the proportions of these subpopulations indicated no correlation to the progress of the learning (Figure 6D and Figure 6-figure supplements 1D).

Earlier version: In the 18F-FDG PET neuroimaging study, the brain activity in the aDLS reached a peak at the middle stage of the acquisition phase of auditory discrimination (Figure 2I). In the multi-unit electrophysiological recording experiment, the firing activity of aDLS neuron subpopulations related to the behavioral outcome showed no significant differences among the three stages (Figure 5E), while the proportions of these subpopulations gradually reduced through the progress of learning stages (Figure 5F). The extent of firing activity and length of firing period of other subpopulations showing sustained activation after the reward appeared to show learning-dependent decrease (Figure 6B and 6C), although the proportions of these subpopulations indicated no correlation to the progress of the learning (Figure 6D). The patterns of the temporal changes in the brain activity of the striatal subregions during the learning stages did not match completely the time variation in the property and proportion of specific event-related neurons. In our electrophysiological analysis, we identified well-isolated neurons from striatal subregions across the behavioral task, focusing on putative MSNs (Figures S4E–S4G). Based on the combinatorial pattern of the tone instruction cue and lever press, we categorized the electrophysiological data into the four trials, including the HR, LL. LR, and HL. We identified HR or LL type neurons showing significant changes in the firing rate related to specific events compared to the baseline firing rate. These neurons were further divided into two groups with increased or decreased activity relative to the baseline firing (Figures S5A and S5B). In the present study, we focused on event-related neurons with the increased activity. Because the analysis was limited to neuronal subpopulations related to specific events with the increased activity, we consider that it is difficult to explain fully dynamic shifts in the brain activity of the striatal subregions dependent on the progress of learning by the time variation of firing activity of individual event-related neurons. The activity of other subpopulations in the striatum may be involved in shifts in brain activity during the learning processes. In addition, recent studies have reported that the activity of glial cells influences the uptake of 18F-FDG28 and that these cells regulate spike timing-dependent plasticity45. Changes in glial cellular activity, through the control of synaptic plasticity, may partly contribute to the pattern formation of learning-dependent shift of brain activity.

Revised version: The patterns of the temporal changes in the brain activity of the striatal subregions during the learning stages did not match completely the time variation in the property and proportion of specific event-related neurons. In our electrophysiological analysis, we identified well-isolated neurons from striatal subregions across the behavioral task, focusing on putative MSNs (Figure 5– figure supplement 2). Based on the combinatorial pattern of the tone instruction cue and lever press, we categorized the electrophysiological data into the four trials, including the HR, LL, HL, and LR. We identified HR or LL type neurons showing significant changes in the firing rate related to specific events compared to the baseline firing rate. These neurons were further divided into two groups with increased or decreased activity relative to the baseline firing (Figure 5–figure supplement 3). In the present study, we focused on event-related neurons with the increased activity. Because the analysis was limited to neuronal subpopulations related to specific events with the increased activity, we consider that it is difficult to explain fully dynamic shifts in the brain activity of the striatal subregions dependent on the progress of learning by the time variation of firing activity of individual event-related neurons. The activity of other subpopulations in the striatum may be involved in the shifts in brain activity during the learning processes. In addition, recent studies have reported that the activity of glial cells influences the uptake of 18F-FDG and that these cells regulate spike timing-dependent plasticity (Valtcheva and Venance, 2016). Changes in glial cellular activity, through the control of synaptic plasticity, may partly contribute to the pattern formation of learning-dependent shift of brain activity.

Materials and methods

Animals

Earlier version: Animal care and handling procedures were carried out in accordance with the guidelines established by RIKEN Center for Biosystems Dynamics Research, Fukushima Medical University, and Osaka City University. All procedures were approved by their Institutional Animal Care and Use Committees. All efforts were made to minimize the number of animals used and their suffering throughout the course of the experiments. Male Long-Evans rats (8–13 weeks old, Institute for Animal Reproduction (Ibaraki, Japan)) were used for the present study. The rats were maintained at 22 ± 2 °C and 60% humidity in a 12-h light/12-h dark cycle, and food and water were continuously available unless stated otherwise.

Revised version: Animal care and handling procedures were carried out in accordance with the guidelines established by RIKEN Center for Biosystems Dynamics Research, Fukushima Medical University, and Osaka City University. All procedures were approved by their Institutional Animal Care and Use Committees. All efforts were made to minimize the number of animals used and their suffering throughout the course of the experiments. Male Long-Evans rats (8–13 weeks old, Institute for Animal Reproduction, Ibaraki, Japan) were used for the present study. The rats were kept at 22 ± 2 °C and 60% humidity in a 12-h light/12-h dark cycle, and were continuously available for laboratory chow (CE-2, CLEA Japan, Fukuoka, Japan) and water. Three days before the behavioral training, the rats were subjected to dietary restriction to maintain their body weight at ∼85% of the free feeding weight or water restriction in which they received free access to water for 6 to 9 h each day.

Surgery

Earlier version: The rats underwent a stereotaxic surgery under isoflurane anesthesia (1–5%). For excitotoxic lesion of striatal neurons, IBO was dissolved in PBS (0.8 mg/mL, Wako Pure Chemical), and IBO solution (0.3 μL per site) was injected bilaterally with the coordinates (mm) from bregma and dura for the aDLS (AP +1.6, ML ±3.3, DV -3.9), pVLS (AP -1.6, ML ±4.6, and DV -5.6), or DMS (AP -1.6, ML ±1.4, and DV -4.8) according to the atlas of rat brain46. Injection was performed through a glass pipette (diameter: 60 μm) at a constant velocity of 0.1 mL/min with the microinfusion pump (EICOM). For transient inactivation of striatal neurons, two guide cannulae (length: 10 mm, outer diameter: 26 gauge) were bilaterally placed on the skull and fixed using four screws and dental cement. The chip of the cannula was placed at 1 mm above the target coordinates. A solution containing MUS (Sigma-Aldrich) in saline (0.1 μg/μL, 0.25 μL per site) was injected into the target coordinates through the internal cannula (outer diameter: 30 gauge) at a constant velocity of 0.25 μL/min. For anterograde and retrograde tracing, a solution containing cholera toxin subunit B (CTb; 1.0 mg/mL) conjugated to Alexa 488 or Alexa 555 (Invitrogen) as a bidirectional tracer was injected into the aDLS and pVLS (0.3 μL/site), respectively. The coordinates for the injection into striatal subregions were the same as described above. For the electrophysiological experiments, surgery was performed as described previously47. A 64-channel silicon probe (Buzsaki64spl; NeuroNexus) attaching to a 3D-printed microdrive with a movable screw was chronically implanted with the coordinates (mm) from the bregma and dura into the left aDLS (AP +1.6, ML -3.3, and DV -2.5) and pVLS (AP -1.6, ML -4.6, and DV -3.6). A faraday cage composed of copper mesh was secured on the skull with dental cement to reduce electrical noise and avoid damage to the implants. The rats were allowed to recover for 1 week after the surgeries.

Revised version: Rats underwent a stereotaxic surgery under isoflurane anesthesia (1–5%). For excitotoxic lesion of striatal neurons, a solution containing IBO (533-54971, FUJIFILM Wako Pure Chemical, Osaka, Japan) in PBS (0.8 mg/mL, 0.3 μL per site) was injected bilaterally with the coordinates (mm) from bregma and dura for the aDLS (AP +1.6, ML ±3.3, DV −3.9), pVLS (AP −1.6, ML ±4.6, and DV - 5.6), or DMS (AP −1.6, ML ±1.4, and DV -4.8) according to the atlas of rat brain (Paxinos and Watson, 1998). Injection was performed through a thin-wall borosilicate glass pipette (diameter: 60 μm, B100-75-10, Sutter Instrument, Novato, CA) at a constant velocity of 0.1 mL/min with the microinfusion pump (ESP-32, EICOM, Kyoto, Japan). For transient inactivation of striatal neurons, two guide cannulae (length: 10 mm, outer diameter: 26 gauge, C315G, P1 Technologies, Roanoke, VA) were bilaterally placed on the skull and fixed using four screws and dental cement. The chip of the cannula was placed at 1 mm above the target coordinates. A solution containing MUS (M1523, Sigma-Aldrich, St. Louis, MO) in saline (0.1 μg/μL, 0.25 μL per site) was injected into the target coordinates through the internal cannula (outer diameter: 30 gauge, C315I, P1 Technologies) at a constant velocity of 0.25 μL/min. For the electrophysiological experiments, surgery was performed as described previously (Kitanishi et al., 2021). A 64-channel silicon probe (Buzsaki64spl, NeuroNexus, Ann Arbor, MI) attaching to a 3D-printed microdrive with a movable screw was chronically implanted with the coordinates (mm) from the bregma and dura into the left aDLS (AP +1.6, ML −3.3, and DV −2.5) and pVLS (AP −1.6, ML -4.6, and DV −3.6). A faraday cage composed of copper mesh was secured on the skull with dental cement to reduce electrical noise and avoid damage to the implants. For anterograde and retrograde tracing, a solution containing CTb conjugated to Alexa 488 (1.0 mg/mL, C34775, Invitrogen, Waltham, MA) or CTb conjugated to Alexa 555 (1.0 mg/mL, C34776, Invitrogen) was injected into the aDLS and pVLS (0.3 μL/site), respectively. The coordinates for the injection into striatal subregions were the same as described above. The rats were allowed to recover for 1 week after the surgeries.

Behavioral analysis

Earlier version: Behavioral tasks were conducted in operant conditioning chambers (30.5 × 24.1 × 29.2 cm; ENV-007CT; MED Associates) equipped with two retractable levers on both sides of a reward port in the front panel (ENV-203M-45, ENV-652AM, and ENV-201A; MED Associates). A multi-tone generator (ENV-223; MED Associates) and room light were mounted on the top of the center and corner panels on the rear side, respectively. For the electrophysical experiment, a 3D printed reward port was used to detect the licking and to prevent contact with the head attachment and front panel. The behavioral experiment in intact rats was performed using grain pellets (F0021, Bioserv) as a food reward. To prevent spillover of radioactivity into the temporal cortex as a consequence of excessive 18F-FDG accumulation into the muscles by their activity to eat the pellets48,49, the PET imaging experiments were carried out using 0.02% saccharine solution (B0131, Tokyo Chemical Industry) as a liquid reward. The electrophysiological experiment was also carried out using the saccharine solution. The pharmacological experiment was performed using the pellets.

Revised version: The behavioral tasks were conducted in operant conditioning chambers (30.5 × 24.1 × 29.2 cm; ENV-007, MED Associates, Fairfax, VT) equipped with two retractable levers on both sides of a reward port in the front panel (ENV-203M-45, ENV-112CM, and ENV-201A, MED Associates). A multi-tone generator (ENV-223, MED Associates) and room light were mounted on the top of the center and corner panels on the rear side, respectively. For the electrophysiological experiment, a 3D printed reward port was used to detect the licking and to prevent contact with the head attachment and front panel. The behavioral experiment was performed using reinforcement pellets (grain-based precision pellets; F0021, Bio-serv, Flemington, NJ) as a food reward. To prevent spillover of radioactivity into the temporal cortex as a consequence of excessive 18F-FDG accumulation into the muscles by their activity to eat the pellets (Endepols et al., 2010; Marx et al., 2012). the PET imaging experiment was carried out using 0.02% saccharin solution (B0131, Tokyo Chemical Industry, Tokyo, Japan) as a liquid reward. The electrophysiological experiment was also carried out by using the saccharin solution.

Earlier version: The rats were subjected to dietary restrictions of 12 g/day (CE-2, CLEA Japan) and water access was limited to 6 to 9 h in accordance with the reward used in the experiments for 3 days before the beginning of the training schedule, and the body weights of the rats were maintained at ∼85% of their normal weights. The training schedule consisted of four continuous steps. In the first step (magazine training), the rats were habituated to the operant conditioning chambers, and the reward was presented with the click every 20 s. Each daily session was started by illumination of the room light, and continued for 30 min. The sessions were conducted for 3 consecutive days. In the second step (shaping), only one lever on the left or right side was inserted. When the rats pressed the lever, they earned a reward accompanied by a click. Each daily session was started by illumination of the light, and lasted 30 min or until the criteria, in which the rats received 90 rewards. When the number of lever presses on either side reached the criteria, the lever was inserted in the opposite side on the next day. When the rats achieved the criteria on both sides for 2 or 3 days, they moved to the following step. In the third step (single lever press task), each trial was started by illumination of the light and insertion of one lever on either the left or right side in a pseudorandom manner with equal frequencies on both sides for every set of four trials. When the rats pressed the lever, the reward was presented with the sound, and the lever was immediately retracted. The room light turned off 4 s after the lever press. When they did not press the lever within 10 s after the start, the trial was terminated when the light turned off. The trial was repeated every 20 s. Each daily session continued for 30 min, and the sessions were conducted for 4 consecutive days. The number of lever presses in each session was counted. In the final step (two-alternative auditory discrimination task), each trial was started by the presentation of a high (10 kHz) or low (2 kHz) frequency tone as an instruction cue in a pseudorandom manner. After three seconds, the room light was illuminated, and two retractable levers were inserted at the same time. The rats were required to press the left or right levers in response to the low or high tone cues, respectively. When the rats pressed the correct lever, the reward was presented with the sound, and the levers were immediately retracted. Correct response allowed the rats to access the reward: the retractable spout delivering water was inserted for 1 s for the PET imaging experiment; a pellet was delivered by a food dispenser for the pharmacological experiment; and three drops of water were delivered through a solenoid valve for the electrophysiological experiment. The room light turned off 4 s after the correct response. On the other hand, error response resulted in the retraction of levers and the turning off of the room light. No response within 10 s after the lever insertion was recorded as an omission response. The success rate in each session was calculated by dividing the number of correct response trials with the number of lever press trials minus the number of omission trials. The trial was repeated every 20 s. Each daily session continued for 60 min, and the sessions were conducted for 14 or 24 days according to the experiments.

Revised version: The training schedule consisted of four continuous steps. In the first step (magazine training), rats were habituated to the operant conditioning chambers, and the reward was presented with the click every 20 s. Each daily session was started by illumination of the room light, and continued for 30 min. The sessions were conducted for 3 consecutive days. In the second step (shaping), only one lever on the left or right side was inserted. When the rats pressed the lever, they earned a reward accompanied by a click. Each daily session was started by illumination of the light, and lasted 30 min or until the criteria, in which the rats received 90 rewards. When the number of lever presses on either side reached the criteria, the lever was inserted in the opposite side on the next day. When the rats achieved the criteria on both sides for 2 or 3 days, they moved to the following step. In the third step (single lever press task), each trial was started by illumination of the light and insertion of one lever on either the left or right side in a pseudorandom manner with equal frequencies on both sides for every set of four trials. When the rats pressed the lever, the reward was presented with the sound, and the lever was immediately retracted. The room light turned off 4 s after the lever press. When they did not press the lever within 10 s after the start, the trial was terminated when the light turned off. The trial was repeated every 20 s. Each daily session continued for 30 min, and the sessions were conducted for 4 consecutive days. The number of lever presses in each session was counted. In the final step (two-alternative auditory discrimination task), each trial was started by the presentation of a high (10 kHz) or low (2 kHz) frequency tone as an instruction cue in a pseudorandom manner. After 3 s, the room light was illuminated, and two retractable levers were inserted at the same time. The rats were required to press the left or right levers in response to the low or high tone cues, respectively. When the rats pressed the correct lever, the reward was presented with the sound, and the levers were immediately retracted. Correct response allowed the rats to access the reward: the retractable spout delivering water was inserted for 1 s for the PET imaging experiment; a pellet was delivered by a food dispenser for the pharmacological experiment; and three drops of water were delivered through a solenoid valve for the electrophysiological experiment. The room light turned off 4 s after the correct response. On the other hand, error response resulted in the retraction of levers and the turning off of the room light.

Added in the revised version: No response within 10 s after the lever insertion was recorded as an omission response. The success rate in each session was calculated by dividing the number of correct response trials with the number of lever press trials minus the number of omission trials. The trial was repeated every 20 s. Each daily session continued for 60 min, and the sessions were conducted for 14 or 24 days according to the experiments.

To assess the behavioral strategies in the auditory discrimination task, a trial-by-trial analysis was conducted; the win-shift-win strategy was defined as a correct response, in which, after a correct response during the previous trial, the rats press the opposite lever during the current trial in response to the shift of the tone instruction cue; the lose-shift-lose strategy was defined as an error response, in which, after an error response during the previous trial, the rats press the opposite lever during the current trial despite the shift of the instruction cue. The proportion of both strategies was calculated by dividing the number of each strategy with the number of lever press trials minus the number of omission trials.

For the transient inactivation experiment, MUS injections were bilaterally performed 20 min before start of the session. For habituation, all rats were injected with saline bilaterally into the corresponding striatal subregions on the last day (Day 4) of the single lever press task. During the auditory discrimination task, the first injection was conducted on Day 2 of the task. The second and third injections were conducted the next day, after the success rate had exceeded 60% and 80%, respectively. Each session continued for 15 min.

Earlier version: For the electrophysiological experiments, a delay period (0.5 ± 0.2 s) was added in the auditory discrimination task between the lever press and reward sound to distinguish the motor- and reward-related responses of neurons. When the rats pressed the lever, the lever was immediately retracted. After the delay period, the reward was presented with a sound. The room light turned off 4 s after the lever press. The trial was repeated every 20 s, and each daily session continued for 60 min.

Revised version: For the electrophysiological experiments, a delay period (0.5 ± 0.2 s) was added in the auditory discrimination task between the lever press and reward sound to distinguish the motor- and reward-related responses of neurons. When the rats pressed the lever, it was immediately retracted. After the delay period, the reward was presented with a sound. The room light turned off 4 s after the lever press. The trial was repeated every 20 s, and each daily session continued for 60 min.

PET imaging analysis

Earlier version: An 18F-FDG-PET scan was performed with a microPET Focus220 (Siemens Medical Solutions, Knoxville) designed for high-resolution imaging of small animals (spatial resolution of 1.4 mm in full width at half maximum at the center of the field of view) as described previously26. Serial 18F-FDG-PET scans were performed on the last day of the single lever press task and on Days 2, 6, 10, and 24 of auditory discrimination task. To administer the 18F-FDG under the awake rats, an indwelling catheter was attached to the tail under anesthesia with 5% induction and 2.0–2.5% maintenance of isoflurane (median [interquartile range]: 6 [5–8] min). The recovery period was set to at least 2 hours until the beginning of the behavioral session, with the aim of minimizing the impact of anesthesia as much as possible. A solution of 18F-FDG (ca. 74 MBq/0.4 mL) was intravenously injected just before the beginning of each behavioral session through the indwelling catheter. After the 30-min session, the rats were returned to their home cage, and then a 30-min static PET scan was performed under anesthesia with 2.0–2.5% isoflurane 55-min after the 18F-FDG injection. During the PET scan, the body temperature of the anesthetized rats was maintained at 37°C using a small animal warmer (BWT-100A, Bio Research Center). The acquired emission data were sorted into a single sinogram and reconstructed by a standard 2D filtered back projection (FBP) with a ramp filter and cutoff frequency at 0.5 cycles per pixel or by a statistical maximum a posteriori probability algorithm (MAP), 12 iterations with point spread function effect.

Revised version: An 18F-FDG-PET scan was performed with a microPET Focus220 (Siemens Medical Solutions, Knoxville, TN) designed for high-resolution imaging of small animals (spatial resolution of 1.4 mm in full width at half maximum at the center of the field of view) as described previously (Cui et al., 2015). Serial 18F-FDG-PET scans were performed on the last day of the single lever press task and on Days 2, 6, 10, and 24 of auditory discrimination task. To administer 18F-FDG under the awake condition, rats were anesthetized through an indwelling catheter attached to their tail with 5% induction and 2.0–2.5% maintenance of isoflurane (median [interquartile range]: 6 [5–8] min). The recovery period was set to at least 2 h until the beginning of the behavioral session, with the aim of minimizing the impact of anesthesia as much as possible. A solution of 18F-FDG (ca. 74 MBq/0.4 mL) was intravenously injected just before the beginning of each behavioral session through the indwelling catheter. After the 30-min session, the rats were returned to their home cage, and then a 30-min static PET scan was performed under anesthesia with 2.0–2.5% isoflurane 55 min after the 18F-FDG injection. During the PET scan, the body temperature of the anesthetized rats was maintained at 37°C using a small animal warmer (BWT-100A, Bio Research Center, Nagoya, Japan). The acquired emission data were sorted into a single sinogram and reconstructed by a standard 2D filtered back projection (FBP) with a ramp filter and cutoff frequency at 0.5 cycles per pixel or by a statistical maximum a posteriori probability algorithm (MAP), 12 iterations with point spread function effect.

Earlier version: Voxel-based analysis of PET images was conducted according to the method described previously26. Briefly, individual MAP-reconstructed FDG (MAP-FDG) images were aligned to a FDG template image of Long Evans rats using PMOD software package (version 3.2, PMOD Technologies), and then the aligned images were transformed into the space of a T2 weighted MRI templates for the Long-Evans rats (https://www.nitrc.org/projects/tpm_rat/). The transformation parameters obtained from individual MAP-FDG images were used to match the individual FBP-reconstructed FDG (FBP-FDG) images with the MRI templates. The voxel size of the template was magnified by 10 to a size similar to that of a human brain. All FBP-FDG images were resampled with a voxel size of 1.2 × 1.2 × 1.2 mm and spatially smoothed through an isotropic Gaussian kernel (6-mm FWHM). Voxel-based statistical analysis was performed by using the Statistical Parametric Mapping (SPM) software (version 8; https://www.fil.ion.ucl.ac.uk/spm/). For visualization of the results, SPM t-statistic values were overlayed on MRI templates to define the voxels showing significantly increased or decreased brain activity, using a high threshold (p = 0.001, uncorrected) and an extent threshold (with the output values from the voxel-based analysis). The voxel of interest (VOI) was determined ass voxels with more than 85% of the maximum T-value within activated areas obtained by voxel-based statistical comparison of FBP-FDG images. The average regional 18F-FDG uptake was calculated on the basis of the values of the VOI in FBP-FDG images.

Revised version: Voxel-based analysis of PET images was conducted according to the method described previously (Cui et al., 2015). Briefly, individual MAP-reconstructed FDG images were aligned to a FDG template image of Long Evans rats using PMOD software package (version 3.2; PMOD Technologies, Fällanden, Switzerland; https://www.pmod.com/web/) and then the aligned images were transformed into the space of a T2 weighted MRI templates for the Long-Evans rat brain (Neuroimaging Tools and Resources Collaboratory; https://www.nitrc.org/projects/tpm_rat/). The transformation parameters obtained from individual MAP-reconstructed FDG images were used to match the individual FBP-reconstructed FDG images with the MRI templates. The voxel size of the template was magnified by 10 to a size similar to that of a human brain. All FBP-reconstructed FDG images were resampled with a voxel size of 1.2 × 1.2 × 1.2 mm and spatially smoothed through an isotropic Gaussian kernel (6-mm FWHM). Voxel-based statistical analysis was performed by using the Statistical Parametric Mapping (SPM) software (version 8; Wellcome Department of Imaging Neuroscience, London, UK; https://www.fil.ion.ucl.ac.uk/spm/). For visualization of the results, SPM t-statistic values were overlayed on MRI templates to define the voxels showing significantly increased or decreased brain activity, using a high threshold (p = 0.001, uncorrected) and an extent threshold (with the output values from the voxel-based analysis). The voxel of interest was determined ass voxels with more than 85% of the maximum T-value within activated areas obtained by voxel-based statistical comparison of the FBP-FDG images. The average regional 18F-FDG uptake was calculated on the basis of values of the voxel of interest in the images.

Electrophysiological recording

Removed in the revised version: Extracellular multi-unit recording was conducted at a sampling rate of 20 kHz from the striatum of freely moving rats by using a 512-ch acquisition board (Open Ephys) via 64-ch recording headstages (C3325; Intan Technologies). The tips of the probe were lowered by turning the screw (> 140 µm/day) until it reached the target region (the aDLS and pVLS) during the single lever press task. The tips were sometimes lowered by the same procedure (140 to 280 µm/day) after the daily recording session to improve cell yield.

Earlier version: Spike sorting was performed automatically by the Kilosort software (https://github.com/cortex-lab/KiloSort). Clear noise clusters were removed in the first manual sorting using the phy GUI (https://github.com/cortex-lab/phy). The remaining clusters were then re-clustered by using the Klusta Kwik (https://klusta.readthedocs.io/en/latest/), and the clusters were adjusted in the second manual sorting by using the phy GUI again. The sorted clusters satisfying all the following criteria were used for further analysis: L-ratio < 0.05, isolation distance > 15, and ISI index < 0.250,51. Well-isolated single neurons (n = 1,485; aDLS, n = 1,062; pVLS, n = 423) were classified as putative medium spiny neurons (MSNs), fast spiking interneurons (FSIs), tonically active neurons (TANs), and unclassified interneurons (UIs). Putative MSNs and TANs were separated from FSIs and UIs by trough-to-peak duration of mean spike waveform (> 0.6 ms) according to previous studies, with slight modifications52,53. Putative FSIs were then separated from UIs by utilizing the proportion of time associated with long inter-spike intervals. This was achieved by summing the inter-spike intervals longer than 2 s and subsequently dividing the resulting sum by the total recording time (PropISIs > 2 s)54. Neurons with values of PropISIs > 2 s less than 0.4 were classified as putative FSIs, while neurons with PropISIs > 2 s greater than 0.4 were classified as putative UIs. The putative MSNs and TANs were separated by measuring the post-spike suppression54. We measured the length of suppressed time during which the firing rate of a neuron was suppressed following a spike by counting the number of 1-ms bins in its autocorrelation function until the rate of the neuron was equal to or greater than its averaged rate over the 600-ms to 900-ms autocorrelation bins. Putative TANs were separated by post-spike suppression of less than 40 ms, and the remaining neurons were classified as putative MSNs. Because the numbers of FSIs, TANs, and UIs were low (79 out of 1,485, 5.3%), they were excluded from further analysis.

Revised version: Extracellular multi-unit recording was conducted at a sampling rate of 20 kHz from the striatum of freely moving rats by using a 512-ch acquisition board (Open Ephys, Atlanta, GA) via 64-ch recording headstages (C3325, Intan Technologies, Los Angeles, CA). The tips of the probe were lowered by turning the screw (> 140 µm/day) until it reached the target regions (the aDLS and pVLS) during the single lever press task. The tips were sometimes lowered by the same procedure (140 to 280 µm/day) after the daily recording session to improve cell yield. Spike sorting was performed automatically by the Kilosort2 software (Stringer et al., 2015; Marius Pachitariu, Ashburn, VA; https://www.github.com/MouseLand/Kilosort2). Clear noise clusters were removed in the first manual sorting using the phy software (cortex-lab, London, UK; https://github.com/cortex-lab/phy). The remaining clusters were then re-clustered by using the KlustaKwik2 software (Kadir et al., 2014; https://github.com/kwikteam/klustakwik2), and the clusters were adjusted in the second manual sorting by using the phy again. The sorted clusters satisfying all the following criteria were used for further analysis: L-ratio < 0.05, isolation distance > 15, and inter-spike interval (ISI) index < 0.2 (Fee et al., 1996; Schmitzer-Torbert et al., 2005). Well-isolated single neurons (n = 1,485; aDLS, n = 1,062; pVLS, n = 423) were classified as putative MSNs, fast spiking interneurons (FSIs), tonically active neurons (TANs), and unclassified interneurons (UIs). Putative MSNs and TANs were separated from FSIs and UIs by trough-to-peak duration of mean spike waveform (> 0.6 ms) according to previous studies with slight modifications (Bakhurin et al., 2016; Kim et al., 2019). Putative FSIs were then separated from UIs by utilizing the proportion of time associated with long ISIs. This was achieved by summing the ISIs longer than 2 s and subsequently dividing the resulting sum by the total recording time (PropISIs > 2 s) (Peters et al., 2021). Neurons with values of PropISIs > 2 s less than 0.4 were classified as putative FSIs, while neurons with PropISIs > 2 s greater than 0.4 were classified as putative UIs. The putative MSNs and TANs were separated by measuring the post-spike suppression (Peters et al., 2021). We measured the length of suppressed time during which the firing rate of a neuron was suppressed following a spike by counting the number of 1-ms bins in its autocorrelation function until the rate of the neuron was equal to or greater than its averaged rate over the 600-ms to 900-ms autocorrelation bins. Putative TANs were separated by post-spike suppression of less than 40 ms, and the remaining neurons were classified as putative MSNs.

Earlier version: To divide the three stages during the acquisition phase, the sigmoidal function was fitted using the following four-parameter logistic equation: Where α is the inflection point, β is a variable for slope, and λ and γ are parameters related to the lower asymptote (γ) and upper asymptote (λ) of the learning curve. The function was used to evaluate the day when the success rate reached 20% (t20%) or 90% (t90%) of the maximum rate change defined by the difference between the γ and λ. Three stages were defined according to the following criteria: early stage (up to 3 days before the middle stage); middle stage (3 days including before and after the day of t20%); and late stage (3 days including the day of t90% and before). One rat learned the discrimination faster, and the success rate reached t90% 2 days after the day of t20%. One day before the day of t90% was assigned to the late stage. For the following analysis, spike data in each striatal subregion were separated into the early, middle, and late stages (n = 295, 364 and 360, respectively, for the aDLS; and n = 87, 110 and 190, respectively, for the pVLS).

Revised version: To divide the three stages during the acquisition phase, the sigmoidal function was fitted using the following four-parameter logistic equation: Where α is the inflection point, β is a variable for slope, and λ and γ are parameters related to the lower asymptote and upper asymptote of the learning curve, respectively. The function was used to evaluate the day when the success rate reached 20% (t20%) or 90% (t90%) of the maximum rate change defined by the difference between the γ and λ. Three stages were defined according to the following criteria: early stage (up to 3 days before the middle stage); middle stage (3 days including before and after the day of t20%); and late stage (3 days including the day of t90% and before). One rat learned the discrimination faster, and the success rate reached t90% 2 days after the day of t20%. One day before the day of t90% was assigned to the late stage. For the following analysis, spike data in each striatal subregion were separated into the early, middle, and late stages (n = 295, 364 and 360, respectively, for the aDLS; and n = 87, 110 and 190, respectively, for the pVLS).

Earlier version: To determine the neurons showing significant changes in firing rate related to the task events, electrophysiological data were grouped into four trial types based on the instruction cue (high or low tone) and the choice (right or left lever); HR, high tone and right lever press; HL, high tone and left lever press; LR, low tone and right lever press; and LL, low tone and left lever press. Neurons with three or more trial data were analyzed. We identified the HR or LL type neurons, detecting significant increases or decreases of the firing rate by applying the Wilcoxon rank-sum test (p < 0.01) to trial-by-trial spike number between the baseline period (BL, -1,500 to -500 ms or -1,100 to -500 ms before presentation of the instruction cue) and any of four peri-event periods (cue onset, 0 to +1,000 ms after the onset of the instruction cue; choice response, -700 to +300 ms around the time of lever press; reward sound, -300 to 300 ms around the time of sound onset instructing the reward delivery; and first licking, 0 to +1,000 ms after the time of first licking after the choice response). Any spikes around the events were averaged in 50-ms bins. The event-related activity was shown as averaged firing rate aligned based on the timing of an event for each trial. The auROC values were calculated by using the spike counts between the HR or LR trials and any of the other three trials. Bins showing significant differences in the auROC data were determined by applying the Wilcoxon rank-sum test (p < 0.05) to the distribution of 200 ms bin-by-bin auROC values and 0.5.

Revised version: To determine the neurons showing significant changes in firing rate related to the task events, electrophysiological data were grouped into four trial types based on the instruction cue (high or low tone) and the choice (right or left lever); HR, high tone and right lever press; LL, low tone and left lever press; HL, high tone and left lever press; and LR, low tone and right lever press. Neurons with three or more trial data were analyzed. We identified the HR or LL type neurons, detecting significant increases or decreases of the firing rate by applying the Wilcoxon rank-sum test (p < 0.01) to trial-by-trial spike number between the baseline period (BL, −1,500 to −500 ms or −1,100 to −500 ms before presentation of the instruction cue) and any of four peri-event periods (cue onset, 0 to +1,000 ms after the onset of the instruction cue; choice response, -700 to +300 ms around the time of lever press; reward sound, −300 to 300 ms around the time of sound onset instructing the reward delivery; and first licking, 0 to +1,000 ms after the time of first licking after the choice response). Any spikes around the events were averaged in 50-ms bins. The event-related activity was shown as averaged firing rate aligned based on the timing of an event for each trial.

Added in the revised version: The auROC values were calculated by using the spike counts between one trial and either of the other three trials. Bins showing significant differences in the auROC data were determined by applying the Wilcoxon rank-sum test (p < 0.05) to the distribution of 200 ms bin-by-bin auROC values and 0.5.

Earlier version: To compare the number of neurons showing event-related activity among different stages, the number of HR or LL type neurons were calculated by matching the trial number at each stage. Twenty trials were randomly selected for a neuron, and neurons with less than 20 trials were removed from the analysis. The neurons with event-related activity were identified as described above. The proportion (percentage) of neurons relative to the total number of neurons was calculated at each individual stage. The same analysis was repeated 30 times and the distribution of the proportion of neurons with event-related activity at the stage was obtained. The cumulative probability against the proportion was calculated using cumulative distribution function (cdfplot, MATLAB code from Mathworks).

Revised version: To compare the number of neurons showing event-related activity among different stages, the number of HR or LL type neurons were calculated by matching the trial number at each stage. Twenty trials were randomly selected for a neuron, and neurons with less than 20 trials were removed from the analysis. The neurons with event-related activity were identified as described above. The proportion (percentage) of neurons relative to the total number of neurons was calculated at each individual stage. The same analysis was repeated 30 times and the distribution of the proportion of neurons with event-related activity at the stage was obtained. The cumulative probability against the proportion was calculated using cumulative distribution function (cdfplot; MATLAB, MathWorks, Natick, MA).

Earlier version: Licking responses were detected with an infrared photobeam sensor placed in the 3D printed reward port, and generated timestamps for the onset of each spout contact according to the analog voltage signal. The licking responses corresponding to four kinds of trial in the first licking-related neurons were shown as averaged rates (50-ms bins) aligned based on the timing of the choice response. To calculate the correlation coefficient between neuronal activity and licking responses, the spikes and licking for 5 s after the first licking were counted every trial for each first licking-related neuron. Trials in which neurons never responded were removed from both the spike and licking data. The correlation coefficients between the two remaining data were calculated for each neuron (corr, MATLAB code from Mathworks), and neurons showing significant differences were determined by applying the Pearson correlation coefficient (p < 0.01).

Revised version: Licking responses were detected with an infrared photobeam sensor placed in the 3D printed reward port, and generated timestamps for the onset of each spout contact according to the analog voltage signal. The licking responses corresponding to four kinds of trial in the first licking- related neurons were shown as averaged rates (50-ms bins) aligned based on the timing of the choice response. To calculate the correlation coefficient between neuronal activity and licking responses, the spikes and licking for 5 s after the first licking were counted every trial for each first licking-related neuron. Trials in which neurons never responded were removed from both the spike and licking data. The correlation coefficients between the two kinds of data were calculated for each neuron (corr; MATLAB, MathWorks), and neurons showing significant differences were determined by applying the Pearson correlation coefficient (p < 0.01).

Histology

Added in the revised version: After the behavioral experiments, rats were anesthetized with pentobarbital (50 mg/kg body weight) and isoflurane (3%), then transcardially perfused 4% paraformaldehyde in phosphate-buffer (PB).

Earlier version: After the behavioral experiments, the rats were anesthetized with pentobarbital (50 mg/kg body weight) and isoflurane (3%), then transcardially perfused 4% paraformaldehyde in phosphate-buffer (PB). The brains were post-fixed overnight at 4 ℃ and cryoprotect with 10, 20, and 30% sucrose in PB. Brain tissues were cut into 30-µm sections by using a cryostat (Leica). For the lesion experiment, cresyl violet staining or immunostaining were performed to define the lesioned areas in the striatum. The sections were stained with a mouse monoclonal anti-NeuN IgG (1:1000 dilution; Millipore, Cat# MAB377, RRID:AB 2298772), and then with a goat anti-mouse IgG conjugated with Alexa Fluor 488 (1:400 dilution; Thermo Fisher Scientific, Cat# A-11029, RRID:AB 2534088). The sections were counterstained with 0.3 μM 4’,6-diamidino-2-phenylindole (DAPI; Sigma-Aldrich). For transient inactivation experiment, sections were stained by cresyl violet to visualize the positions of guide cannulae. For the electrophysiological experiments, small lesions were made by applying an electrolytic current (3 μA, 10 s) through top and bottom channels of each shank by using the stimulus isolator (A365, World Precision Instrument). Sections through the striatum were stained with a mouse monoclonal anti-TH IgG (1:1,000 dilution; Millipore, Cat# MAB318, RRID:AB 2201528), and then with a goat anti-mouse IgG conjugated with Alexa Fluor 488 (1:400 dilution; Thermo Fisher Scientific, Cat# A-11029, RRID:AB 2534088). The recording sites were verified by the staining with DAPI and red fluorescent Nissl stain solution (1:200 dilution; Thermo Fisher Scientific, Cat# N21482, RRID:AB 2620170). Fluorescence and immunostained signals were visualized under fluorescence microscopes (BZ-X800; Keyence) and confocal microscope (A1R; Nikon).

Revised version: The brains were post-fixed overnight at 4 ℃ and cryoprotect with 10, 20, and 30% sucrose in PB. Brain tissues were cut into 30-µm sections by using a cryostat (CM3050 S, Leica Biosystems, Wetzlar, Germany). To define the lesioned areas in the striatum, sections were immunostained with a mouse anti-NeuN antibody (1:1,000 dilution; MAB377, Millipore, Burlington, MA), and then with a goat anti-mouse IgG antibody conjugated with Alexa Fluor 488 (1:1,000 dilution; A-11029, Thermo Fisher Scientific, Waltham, MA). The sections were counterstained with 0.3 μM 4’,6- diamidino-2-phenylindole (DAPI; 0.1 μg/mL, D1306, Thermo Fisher Scientific). For the transient inactivation experiment, the positions of guide cannulae were confirmed by staining with 0.1% cresyl violet solution (41021, Muto Pure Chemicals, Tokyo, Japan). For the electrophysiological experiments, small lesions were made by applying an electrolytic current (3 μA, 10 s) through top and bottom channels of each shank by using the stimulus isolator (A365, World Precision Instruments, Sarasota, FL). Sections through the striatum were stained with a mouse anti-TH antibody (1:1,000 dilution; MAB318-AF488, Millipore), and then with a goat anti-mouse IgG antibody conjugated with Alexa Fluor 488 (1:1,000 dilution; Jackson Immunoresearch Laboratories, West Grove, PA). The recording sites were verified by staining with DAPI and red fluorescent Nissl stain solution (1:10,000/1:200 dilution; N21482, Thermo Fisher Scientific). For the tracing, sections through the ventral midbrain were stained with a mouse anti-TH antibody (1:1,000 dilution), and then with a donkey anti-mouse IgG antibody conjugated with Alexa Fluor 647 (1:1,000 dilution; A- 31571, Thermo Fisher Scientific). Fluorescence and immunostained signals were visualized under fluorescence microscopes (BZ-X800, Keyence, Osaka, Japan) and confocal microscope (A1R, Nikon, Osaka, Japan).

Statistical analysis

Added in the revised version: Statistical analysis was performed using MATLAB (MathWorks) and SPSS Statistics (version 27; IBM, Armonk, NY). Comparisons of data between two groups were performed using the Student’s t-test (paired), the Wilcoxon signed rank test (paired), and the Wilcoxon rank-sum test (unpaired).

Earlier version: Statistical analysis was performed using MATLAB (MathWorks) and SPSS Statistics (Ver. 27, IBM). Comparisons of data between two groups were performed using the Student’s t-test (paired), the Wilcoxon signed rank test (paired), and the Wilcoxon rank-sum test (unpaired). Comparisons of data among three groups were constructed using a one-way repeated measures ANOVA, two-way repeated measures ANOVA, or the Kruskal-Wallis test (without assuming normal distribution and homogeneity of variances). All statistical tests were two-sided. The Bonferroni post-hoc test and Mann–Whitney U-test with Bonferroni adjustment were carried out as needed to compare means. For boxplots, the data points below the lower (25th percentile) -1.5 interquartile range (IQR) and above the upper (75th percentile) quartiles +1.5 IQR were regarded as outliers, and they are not included in the boxplots.

Revised version: Comparisons of data among three groups were constructed using a one-way repeated measures ANOVA, two-way repeated measures ANOVA, or the Kruskal-Wallis test (without assuming normal distribution and homogeneity of variances). All statistical tests were two-sided. The Bonferroni post-hoc test and Mann–Whitney U-test with Bonferroni adjustment were carried out as needed to compare means. For boxplots, the data points below the lower (25th percentile) −1.5 interquartile range and above the upper (75th percentile) quartiles +1.5 interquartile range were regarded as outliers, and they are not included in the boxplots.